Evidence that humans underestimate body mass in microgravity: kinematic signatures in reaching movements during spaceflight

Author response:

The following is the authors’ response to the original reviews.

General recommendations (from the Reviewing Editor):

The reviewers discussed the revision at length, and all were appreciative of the revisions to the paper. Nonetheless, they agreed that the evidence against alternative hypotheses was not yet decisive, and it may not be possible to provide the evidence needed given the difficulty of acquiring this data. Thus they feel that a more nuanced interpretation of the data and tempering of the conclusions is necessary. These points are described in more detail in the reviewer-specific comments in the Public reviews.

We thank the editor and the reviewers for their constructive discussion. In this revision, we have adopted these recommendations: we have tempered our conclusions and removed binary framing, taking into consideration that other alternative explanations might exist. We have also expanded the Discussion to consider additional potential mechanisms and added corresponding limitations. We also changed the paper title to avoid strong inference; the new title is “Evidence that humans underestimate body mass in microgravity: kinematic signatures in reaching movements during spaceflight”.

Public Reviews:

Reviewer #1 (Public review):

The authors have conducted substantial additional analyses to address the reviewers' comments. However, several key points still require attention. I was unable to see the correspondence between the model predictions and the data in the added quantitative analysis. In the rebuttal letter, the delta peak speed time displays values in the range of [20, 30] ms, whereas the data were negative for the 45{degree sign} direction. Should the reader directly compare panel B of Figure 6 with Figure 1E? The correspondence between the model and the data should be made more apparent in Figure 6. Furthermore, the rebuttal states that a quantitative prediction was not expected, yet it subsequently argues that there was a quantitative match. Overall, this response remains unclear.

We thank the reviewer raising the question about Figure 6B. We would like to clarify that the phrase "quantitative match" in the summary of our previous rebuttal letter was a wording error; in fact, the subsequent detailed responses consistently and correctly described the comparison as qualitative. We apologize for the confusion this may have caused, and address this point below.

First, we have revised the manuscript to clarify this point. We have added the following statement: "We note that these correlations evaluate the directional trend rather than the absolute magnitude of the effects; a precise quantitative match is not expected given the simplifications of the two-joint arm model." in the main text.

Second, we have replaced Figure 6 with a revised version that presents model-predicted Δ values and experimentally observed Δ values side by side, allowing for a more intuitive visual comparison. As shown in the updated figure, the directional trends are broadly consistent amplitude changes and timing shifts are rank-ordered by movement direction in both model and data while the absolute magnitudes do not precisely match. We believe this layout makes the intended comparison more transparent.

As discussed in our previous response, as noted above, a precise quantitative match is not expected given our model's simplifications, and this level of qualitative comparison is consistent with established practice in similar modeling studies (e.g., Gaveau et al., 2016).

Regarding the negative Δ peak speed time at 45°: as shown in our statistical analyses (Figure 4A, Figure 5F), there was no significant timing change at 45°. The negative value reflects a small, non-significant mean difference. The key pattern that timing advance increases for directions associated with higher effective inertia holds for the 90° and 135° directions, which is the directional trend our analysis was designed to capture.

A follow-up question concerns the argument about strategic slowing. The authors argue that this explanation can be rejected because the timing of peak speed should be delayed, contrary to the data. However, there appears to be a sign difference between the model and the data for the 45{degree sign} direction, which means that it was delayed in this case. Did I understand correctly? In that regard, I believe that the hypothesis of strategic slowing cannot yet be firmly rejected and the discussion should more clearly indicate that this argument is based on some, but not all, directions.

I agree with the authors on the importance of the mass underestimation hypothesis, and I am not particularly committed to the strategic slowing explanation, but I do not see a strong argument against it. If the conclusion relies on the sign of the delta peak speed, then the authors' claims are not valid across all directions, and greater caution in the interpretation and discussion is warranted. Regarding the peak acceleration time, I would be hesitant to draw firm conclusions based on differences smaller than 10 ms (Figures R3 and 6D).

The authors state in the rebuttal that the two hypotheses are competing. This is not accurate, as they are not mutually exclusive and could even vary as a function of movement direction. The abstract also claims that the data "refutes" strategic slowing, which I believe is too strong. The main issue is that, based on the authors' revised manuscript, the lack of quantitative agreement between the model and the data for the mass underestimation hypothesis is considered acceptable because a precise quantitative match is not expected, and the predictions overall agree for some (though not all) directions and phases (excluding post-in). That is reasonable, but by the same logic, the small differences between the model prediction and the strategic slowing hypothesis should not be taken as firm evidence against it, as the authors seem to suggest. In practice, I recommend a more transparent and cautious interpretation to avoid giving readers the false impression that the evidence is decisive. The mass underestimation hypothesis is clearly supported, but the remaining aspects are less clear, and several features of the data remain unexplained.

We thank the reviewer for this critical assessment. We acknowledge that our previous framing was too binary, and we agree that strategic slowing and mass underestimation are not mutually exclusive. We would like to clarify our view: we did not find evidence supporting strategic slowing (e.g., slower reaction times, symmetric velocity/acceleration peaks), whereas we did find evidence supporting mass underestimation (asymmetric peaks, unchanged reaction times, more sub movements). This is not a case of rejecting one hypothesis to affirm the other; our data simply do not support one while providing positive evidence for the other. We do not rule out the possibility that both mechanisms could operate together, though we note that our data did not reveal evidence supporting strategic slowing in the current reaching task.

We also agree that the lack of significant timing changes at 45° limits the scope of our argument against strategic slowing in that direction. However, the null result at 45° likewise cannot serve as positive evidence for strategic slowing either. As discussed in our previous revision and in Discussion, this null effect may arise because 45° reaches are predominantly single-joint (evidenced by curvature patterns characteristic), making them less suitable for modeling with a simplified two-link arm model than the 90° and 135° directions.

In line with these considerations, we have made the following revisions to the manuscript:

(1) We have removed binary framing throughout, replacing claims of mutual exclusivity or outright rejection of strategic slowing with more measured language. For example, "refutes" in the abstract has been changed to "These findings provide support for the body mass underestimation hypothesis while being inconsistent with the strategic slowing hypothesis." The two hypotheses are no longer presented as mutually exclusive, and strategic slowing is now characterized as insufficient to fully explain the direction-dependent pattern, rather than ruled out entirely.

(2) We have revised the conclusion. The concluding paragraph no longer presents an either-or outcome. We describe the direction-dependent under-actuation pattern, note that it strongly supports mass underestimation while not being readily explained by a uniform strategic adjustment, and acknowledge that other factors may also contribute. A new limitation paragraph discusses the simplified nature of our model and acknowledges that other neurophysiological and biomechanical factors cannot be excluded.

Reviewer #2 (Public review):

This study explores the underlying causes of the generalized movement slowness observed in astronauts in weightlessness compared to their performance on Earth. The authors argue that this movement slowness stems from an underestimation of mass rather than a deliberate reduction in speed for enhanced stability and safety.

Overall, this is a fascinating and well-written work. The kinematic analysis is thorough and comprehensive. The design of the study is solid, the collected dataset is rare, and the model adds confidence to the proposed conclusions.

Compared to the previous version, the authors have thoroughly addressed my concerns. The model is now clear and well-articulated, and alternative hypotheses have been ruled out convincingly. The paper is improved and suitable for publication in my opinion, making a significant contribution to the field.

Strengths:

Comprehensive analysis of a unique data set of reaching movement in microgravity

Use of a sensible and well-thought experimental approach

State-of-the-art analyses of main kinematic parameter

Computational model simulations of arm reaching to test alternative hypotheses and support the mass underestimation one

This work has no major weakness as it stands, and the discussion provides a fair evaluation of the findings and conclusions.

We thank the reviewer for the supportive feedback, and we are grateful for the earlier comments that helped us improve the manuscript.

Reviewer #3 (Public review):

Summary:

The authors describe an interesting study of arm movements carried out in weightlessness after a prolonged exposure to the so-called microgravity conditions of orbital spaceflight. Subjects performed radial point-to-point motions of the fingertip on a touch pad. The authors note a reduction in movement speed in weightlessness, which they hypothesize could be due to either an overall strategy of lowering movement speed to better accommodate the instability of the body in weightlessness or an underestimation of body mass. They conclude for the latter, mainly based on two effects. One, slowing in weightlessness is greater for movement directions with higher effective mass at the end effector of the arm. Two, they present evidence for increased number of corrective submovements in weightlessness. They contend that this provides conclusive evidence to accept the hypothesis of an underestimation of body mass.

Strengths:

In my opinion, the study provides a valuable contribution, the theoretical aspects are well presented through simulations, the statistical analyses are meticulous, the applicable literature is comprehensively considered and cited and the manuscript is well written.

Weaknesses:

I nevertheless am of the opinion that the interpretation of the observations leaves room for other possible explanations of the observed phenomenon, thus weakening the strength of the arguments.

To strengthen the conclusions, I feel that the following points would need to be addressed:

We thank the reviewer for the insightful critique and constructive suggestions. Following the reviewer's advice, we have re-framed our Introduction and Discussion to present mass underestimation as a plausible mechanism identified by our simplified model, while explicitly acknowledging other potential factors. Below we address each point in detail.

(1) The authors model the movement control through equations that derive the input control variable in terms of the force acting on the hand and treating the arm as a second-order low pass filter (Eq. 13). Underestimation of the mass in the computation of a feedforward command would lead to a lower-than-expected displacement to that command. But it is not clear if and how the authors account for a potential modification of the time constants of the 2nd order system. The CNS does not effectuate movements with pure torque generators. Muscles have elastic properties that depend on their tonic excitation level, reflex feedback and other parameters. Indeed, Fisk et al.* showed variations of movement characteristics consistent with lower muscle tone, lower bandwidth and lower damping ratio in 0g compared to 1g. Could the variations in the response to the initial feedforward command be explained by a misrepresentation of the limbs damping and natural frequency, leading to greater uncertainty to the consequences of the initial command. This would still be an argument for un-adapted feedforward control of the movement, leading to the need for more corrective movements. But it would not necessarily reflect an underestimation of body mass.

*Fisk, J. O. H. N., Lackner, J. R., & DiZio, P. A. U. L. (1993). Gravitoinertial force level influences arm movement control. Journal of neurophysiology, 69(2), 504-511.

While the authors attempt to differentiate their study from previous studies where limb neuromechanical impedance was shown to be modified in weightlessness by emphasizing that in the current study the movements were rapid and the initial movement is "feedforward". But this incorrectly implies that the limb's mechanical response to the motor command is determined only by active feedback mechanisms. In fact:

(a) All commands to the muscle pass through the motor neurons. These neurons receive descending activations related not only to the volitional movement, but also to the dynamic state of the body and the influence of other sensory inputs, including the vestibular system. A decrease in descending influences from the vestibular organs will lower the background sensitivity to all other neural influences on the motor neuron. Thus, the motor neuron may be less sensitive to the other volitional and reflexive synaptic inputs that it may receive.

(b) Muscle tone plays a significant role in determining the force and the time course of the muscle contraction. In a weightless environment, where tonic muscle activity is likely to be reduced, there is the distinct possibility that muscles will react more slowly and with lower amplitude to an otherwise equivalent descending motor command, particularly in the initial moments before spinal reflexes come into play. These, and other neuronal mechanisms could lead to the "under-actuation" effect observed in the current study, without necessarily being reflective of an underestimation of mass per se.

The reviewer raises an important point that the observed underactuation may not necessarily reflect mass underestimation per se. It could also arise from changes in the time constants of the control system, tonic muscle activation levels, vestibular descending inputs, or altered spinal reflex gains. We agree that our simplified model does not capture these neuromuscular factors, and we have made several revisions to address this concern.

In the Discussion (paragraph 4), we have added a new substantive section discussing how reduced tonic muscle activity, diminished vestibular inputs to motor neurons, and altered muscle activation dynamics (Fisk et al., 1993) may contribute to the observed under-actuation independently of mass misestimation. We argue that while these factors likely affect motor output, they would be expected to produce a relatively uniform effect across movement directions, as tonic muscle activation and vestibular descending inputs are not specific to a particular reaching direction. In contrast, the direction-dependent pattern of our results with greater effects for directions involving higher effective mass is more naturally explained by a misrepresentation of inertial properties than by a uniform change in neuromuscular excitability. Nevertheless, we explicitly acknowledge that these mechanisms may act in concert with mass underestimation, and that our current data cannot fully disentangle them.

Additionally, the paragraph discussing proprioceptive mechanisms (paragraph 6 of Discussion) now opens with the conditional framing "If mass underestimation contributes to the observed underactuation," and closes by noting that the same proprioceptive degradation could affect motor output through other pathways such as reducing tonic muscle activation or altering spinal reflex gains independent of any explicit misrepresentation of body mass.

We have also added a new limitation (the fourth in the Limitations section) explicitly acknowledging that our model treats muscles as ideal torque generators and does not capture potential changes in muscle activation dynamics, damping, or reflex gains that may occur in microgravity. Future studies combining detailed musculoskeletal modeling with direct measurements of muscle activation, joint impedance, and trunk kinematics would be needed to distinguish between mass underestimation and other sources of underactuation.

That said, the assumption of relatively preserved muscle properties is partly supported by the available evidence. A systematic review of simulated microgravity studies found that upper limb maximal voluntary contraction remained mostly unchanged for up to 45 days of unloading, and that upper limb muscles declined substantially more slowly than lower limb and trunk muscles (Winnard et al., 2019). A more recent review similarly reported that upper limb muscle outcomes are less affected by microgravity exposure (Bosutti et al., 2025). This is also consistent with our own unpublished observations in Chinese astronauts, which did not indicate an obvious decline in upper limb force output. While these findings do not rule out subtler changes in muscle tone or activation dynamics, they suggest that gross alterations in upper limb neuromuscular capacity are unlikely to be the primary driver of the underactuation we observed.

Refs.

Winnard, A., Scott, J., Waters, N., Vance, M., & Caplan, N. (2019). Effect of time on human muscle outcomes during simulated microgravity exposure without countermeasures—systematic review. Frontiers in physiology, 10, 1046.

Bosutti, A., Ganse, B., Maffiuletti, N. A., Wüst, R. C., Strijkers, G. J., Sanderson, A., & Degens, H. (2025). Microgravity‐induced changes in skeletal muscle and possible countermeasures: What we can learn from bed rest and human space studies. Experimental Physiology.

(2) The subject's body in weightless is much more sensitive to reaction forces in interactions with the environment in the absence of the anchoring effect of gravity pushing the body into the floor and in the absence of anticipatory postural adjustments that typically accompany upper-limb motions in Earth gravity in order to maintain an upright posture. The authors dismiss this possibility because the taikonauts were asked to stabilize their bodies with the contralateral hand. But the authors present no evidence that this was sufficient to maintain the shoulder and trunk at a strictly constant position, as is supposed by the simplified biomechanical model used in their optimal control framework. Indeed, a small backward motion of the shoulder would result in a smaller acceleration of the fingertip and a smaller extent of the initial ballistic motion of the hand with respect to the measurement device (the tablet), consistent with the observations reported in the study. Note that stability of the base might explain why 45º movements were apparently less affected in weightlessness, according to many of the reported analyses, including those related to corrective movements (Fig. 5 B, C, F; Fig. 6D), than the other two directions. If the trunk is being stabilized by the left arm, the same reaction forces on the trunk due to the acceleration of the hand will result in less effective torque on the trunk, given that the reaction forces act with a much smaller moment arm with respect to the left shoulder (the hand movement axis passes approximately through the left shoulder for the 45º target) compared to either the forward or rightward motions of the hand.

The reviewer raises an important point about the potential influence of reaction forces on trunk and shoulder stability in microgravity. We have revised the relevant Discussion paragraph to address this concern more thoroughly.

We would like to clarify that, in addition to stabilizing the body with the left hand grasping a fixed bar, the taikonauts’ feet were also constrained with foot straps, providing multi-point stabilization. Furthermore, the reviewer's trunk displacement hypothesis predicts that the 45° direction should be systematically less affected across all kinematic measures. However, while 45° did not show significant changes in the timing of kinematics peaks, it did show significant changes in movement duration, peak acceleration, and peak speed comparable to the other directions. This dissociation is difficult to reconcile with a uniform trunk displacement artifact, but is consistent with a direction-dependent inertial effect.

We acknowledge that we did not directly measure trunk or shoulder kinematics, highlight that we did our best to provide multi-point stabilization in our setup, and we have added this as a limitation in the revised Discussion.

(3) The above is exacerbated by potential changes in the frictional forces between the fingertip and the tablet. The movements were measured by having the subjects slide their finger on the surface of a touch screen. In weightlessness, the implications of this contact can be expected to be quite different than on the ground. While these forces may be low on Earth, the fact is that we do not know what forces the taikonauts used on orbit. In weightlessness, the taikonauts would need to actively press downward to maintain contact with the screen, while on Earth gravity will do the work. The tangential forces that resist movement due to friction might therefore be different in 0g. . Indeed, given the increased instability of the body and the increased uncertainty of movement direction of the hand, taikonauts may have been induced to apply greater forces against the tablet in order to maintain contact in weightlessness, which would in turn slow the motion of the finger on the table and increase the reaction forces acting on the trunk. This could be particularly relevant given that the effect of friction would interact with the limb in a direction-dependent fashion, given the anisotropy of the equivalent mass at the fingertip evoked by the authors

We agree that in microgravity, taikonauts must actively press on the screen to maintain contact, potentially altering normal forces and thus friction compared to ground conditions. We have acknowledged this point in the revised Discussion. However, we note several reasons why friction is unlikely to be the dominant factor. First, the tablet uses a capacitive touchscreen, which registers touch through changes in electrical capacitance and does not require substantial normal force to maintain contact. Second, typical tangential friction forces during touchscreen interaction range from 0.1 to 0.5 N (Ayyildiz et al., 2018), which are small compared to the 10–15 N required to accelerate the arm during reaching. Third, touchscreen performance has been shown to be largely unaffected during long-duration spaceflight (Holden et al., 2022). Lastly but importantly, the friction hypothesis does not readily account for the direction-specific pattern of effects we observed. While we cannot exclude a contribution of altered friction, particularly in interaction with the direction-dependent effective mass, its magnitude makes it unlikely to account for the observed kinematic changes.

Ref:

Ayyildiz, M., Scaraggi, M., Sirin, O., Basdogan, C., & Persson, B. N. J. (2018). Contact mechanics between the human finger and a touchscreen under electroadhesion. Proceedings of the National Academy of Sciences of the United States of America, 115(50), 12668–12673.

Holden, K., Greene, M., Vincent Cross, E., Sandor, A., Thompson, S., Feiveson, A., & Munson, B. (2023). Effects of long-duration microgravity and gravitational transitions on fine motor skills. Human Factors, 65(6), 1046-1058.

I feel that the authors have done an admirable job of exploring the how to explain the modifications to movement kinematics that they observed on orbit within the constraints of the optimal control theory applied to a simplified model of the human motor system. While I fully appreciate the value of such models to provide insights into question of human sensorimotor behaviour, to draw firm conclusions on what humans are actually experiencing based only on manipulations of the computational model, without testing the model's implicit assumptions and without considering the actual neurophysiological and biomechanical mechanisms, can be misleading. One way to do this could be to examine these questions through extensions to the model used in the simulations (changing activation dynamics of the torque generators, allowing for potential motion backward motion of the shoulder and trunk, etc.). A better solution would be to emulate the physiological and biomechanical conditions on Earth (supporting the arm against gravity to reduce muscle tone, placing the subject on a moveable base that requires that the body be stabilized with the other hand) in order to distinguish the hypothesis of an underestimation of mass vs. other potential sources of under-actuation and other potential effects of weightlessness on the body.

In sum, my opinion is that the authors are relying too much on a theoretical model as a ground truth and thus overstate their conclusions. But to provide a convincing argument that humans truly underestimate mass in weightlessness, they should consider more judiciously the neurophysiology and biomechanics that fall outside the purview of the simplified model that they have chosen. If a more thorough assessment of this nature is not possible, then I would argue that a more measured conclusion of the paper should be 1) that the authors observed modifications to movement kinematics in weightlessness consistent with an under-actuation for the intended motion, 2) that a simplified model of human physiology and biomechanics that incorporates principles of optimal control suggest that the source of this under-actuation might be an underestimation of mass in the computation of an appropriate feedforward motor command, and 3) that other potential neurophysiological or biomechanical effects cannot be excluded due to limitations of the computational model.

We appreciate the reviewer's thoughtful assessment. We fully agree that a simplified computational model should not be treated as ground truth, and that the neurophysiology and biomechanics beyond the computational model must be carefully considered.

As detailed in our responses above, we have substantially revised the Discussion to address each of these concerns—including new discussions of neuromuscular factors, more balanced treatment of trunk stability and friction, conditional framing of the mass underestimation interpretation, and a new limitation on model simplifications. The conclusion has been restructured following the reviewer's recommended framework.

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

If possible and allowed, the authors are strongly encouraged to consider sharing this unique dataset. Making the data publicly available alongside the paper could foster future studies and further accelerate research in this area.

We sincerely thank the reviewer for this suggestion. The ground control data and all analysis code will be made publicly available alongside the Version of Record.

However, unfortunately, the raw in-flight data from the taikonaut cohort cannot be made publicly available due to confidentiality regulations of China's manned space program; access for scientific research requires approval from the China Astronaut Research and Training Center and can be requested through the corresponding author.

Changes from the earlier version to the revised version

Abstract

Earlier version: Astronauts consistently exhibit slower movements in microgravity, even during tasks requiring rapid responses. The sensorimotor mechanisms underlying this general slowing remain debated. Two competing hypotheses have been proposed: either the sensorimotor system adopts a conservative control strategy for safety and postural stability, or the system underestimates body mass due to reduced inputs from proprioceptive receptors. To resolve the debate, we studied twelve taikonauts aboard the China Space Station performing a classical hand-reaching task. Compared to their pre-flight performance and to an age-matched control group, participants showed increased movement durations and altered kinematic profiles in microgravity. Model-based analyses of motor control parameters revealed that these changes stemmed from reduced initial force generation in the feedforward control phase followed by compensatory feedback-based corrections. These findings support the body mass underestimation hypothesis while refuting the strategic slowing hypothesis. Importantly, the sensory estimate of bodily property in microgravity is biased but immune from sensorimotor adaptation, calling for an extension of existing theories of motor learning.

Revised version: Astronauts consistently exhibit slower movements in microgravity, even during tasks requiring rapid responses. The sensorimotor mechanisms underlying this general slowing remain debated. Two hypotheses have been proposed: either the sensorimotor system adopts a conservative control strategy for safety and postural stability, or the system underestimates body mass due to reduced inputs from proprioceptive receptors. To dissociate these opinions, we studied twelve taikonauts aboard the China Space Station performing a classical hand-reaching task. Compared to their pre-flight performance and to an age-matched control group, participants showed increased movement durations and altered kinematic profiles in microgravity. Model-based analyses of motor control parameters revealed that these changes stemmed from reduced initial force generation in the feedforward control phase followed by compensatory feedback-based corrections. These findings providing support for the body mass underestimation hypothesis while being inconsistent with the strategic slowing hypothesis. Importantly, the sensory estimate of bodily property in microgravity is biased but immune from sensorimotor adaptation, calling for an extension of existing theories of motor learning.

Introduction

Successful space exploration fundamentally depends on efficient sensorimotor control under microgravity conditions, where both movement accuracy and speed are critical (Strangman et al., 2014). A consistently observed phenomenon is that astronauts exhibit slower movements during routine activities in space stations (Kubis et al., 1977). This general movement slowing persists even in controlled experimental settings where astronauts are securely fastened without concerns for postural stability or safety. Indeed, fast movements in microgravity can be disadvantageous, potentially destabilizing posture, complicating the arrest of accelerated body segments, and increasing collision risks in confined spacecraft environments (Bock, 1998; Layne et al., 2001; Paloski et al., 1999; Reschke et al., 1998; Tafforin et al., 1989). Yet even when body stability is ensured, well-practiced goal-directed actions consistently show prolonged movement durations across various tasks, from hand pointing to visual targets (Berger et al., 1997; Bock et al., 2001, 2003; Mechtcheriakov et al., 2002) to step tracking with joysticks (Newman & Lathan, 1999; Sangals et al., 1999). This slowing manifests as reduced peak speeds in hand pointing (Mechtcheriakov et al., 2002; Sangals et al., 1999; Weber & Stelzer, 2022). Despite extensive documentation of this phenomenon, the underlying sensorimotor mechanisms remain unclear (Bock et al., 1996b; Weber et al., 2019), presenting a significant challenge for optimizing human performance in space.

Two primary hypotheses have emerged to explain the slowing of goal-directed actions in microgravity. The conservative control hypothesis suggests that the sensorimotor system adopts a generalized strategy prioritizing safety and postural stability over speed (Berger et al., 1997; Bock et al., 2010; Mechtcheriakov et al., 2002). This adaptation may arise from the novel challenges of movement control in microgravity, where even simple actions can propagate unexpected forces through body segments. The persistent application of this conservative strategy, as a generalized adaptation to the novel environment, could explain the movement slowing in microgravity. Alternatively, the body mass underestimation hypothesis builds on observations that humans consistently underestimate the mass of handheld objects in microgravity (Ross & Reschke, 1982), suggesting that similar perceptual biases might extend to the estimation of body segment properties (Bock et al., 1996b). Such misperception would directly impact the feedforward control of movement, potentially leading to systematic under-actuation of initial movements (Elliott et al., 2010, 2017).

Earlier version: Both hypotheses predict reduced peak speed and acceleration in microgravity, but they diverge in their predictions about movement kinematics. First, typical reaching movement exhibits a symmetrical bell-shaped speed profile, which minimizes energy expenditure while maximizing accuracy or smoothness according to optimal control principles (Flash & Hogan, 1985; Todorov, 2004). The conservative control hypothesis suggests a maintained temporal symmetry as longer durations are strategically planned for optimal performance, the peak speed and acceleration are thus delayed in time (Figure 1A). In contrast, the body mass underestimation hypothesis predicts that insufficient initial force generation, resulting from a systematic sensory bias in internal models that predict movement dynamics, would lead to earlier occurrence of peak speed and acceleration (Elliott et al., 2010, 2017). Previously, the timing of these peaks has shown mixed results, with some studies reporting earlier peaks (Sangals et al., 1999), others finding delays (Fowler et al., 2008), and some showing no temporal shifts (Berger et al., 1997; Mechtcheriakov et al., 2002). Second, the hypotheses make distinct predictions about feedback-based corrections during late reaching. The conservative control hypothesis suggests that movements are optimally adapted for the microgravity environment, thus requiring no additional corrections. Instead, the mass underestimation hypothesis predicts that initial underactuation would necessitate subsequent feedback-based corrections to reach the target (Chua & Elliott, 1993; Smeets & Brenner, 1995). These corrections would manifest as additional submovements, contributing to asymmetric speed profiles (Figure 1A). Current evidence for increased feedback-based corrections in microgravity remains inconsistent: while some studies support their presence (Sangals et al., 1999), others show conflicting results (Mechtcheriakov et al., 2002) or opposite effects (Fowler et al., 2008). If present, these corrective submovements should predict movement duration prolongation, serving as crucial evidence for the mass-underestimation account for the general slowing in microgravity.

Revised version: Both hypotheses predict reduced peak speed and acceleration in microgravity, but they differ in their predictions about finer-grained movement kinematics. First, typical reaching movement exhibits a symmetrical bell-shaped speed profile, which minimizes energy expenditure while maximizing accuracy or smoothness according to optimal control principles (Flash & Hogan, 1985; Todorov, 2004). The conservative control hypothesis suggests a maintained temporal symmetry as longer durations are strategically planned for optimal performance, the peak speed and acceleration are thus delayed in time (Figure 1A). In contrast, the body mass underestimation hypothesis predicts that insufficient initial force generation, resulting from a systematic sensory bias in internal models that predict movement dynamics, would lead to earlier occurrence of peak speed and acceleration (Elliott et al., 2010, 2017). Previously, the timing of these peaks has shown mixed results, with some studies reporting earlier peaks (Sangals et al., 1999), others finding delays (Fowler et al., 2008), and some showing no temporal shifts (Berger et al., 1997; Mechtcheriakov et al., 2002). Second, the hypotheses make distinct predictions about feedback-based corrections during late reaching. The conservative control hypothesis suggests that movements would remain well-planned with near-optimal execution, thus requiring minimal corrections. Instead, the mass underestimation hypothesis predicts that initial underactuation would necessitate subsequent feedback-based corrections to reach the target (Chua & Elliott, 1993; Smeets & Brenner, 1995). These corrections would manifest as additional submovements, contributing to asymmetric speed profiles (Figure 1A). Current evidence for increased feedback-based corrections in microgravity remains inconsistent: while some studies support their presence (Sangals et al., 1999), others show conflicting results (Mechtcheriakov et al., 2002) or opposite effects (Fowler et al., 2008). If present, these corrective submovements should predict movement duration prolongation, serving as important evidence favoring the mass-underestimation account for the general slowing in microgravity.

Model simulations revealing characteristic kinematic features of reaching movements.

Earlier version: A) Speed profile predictions under two competing hypotheses. Under normal gravity, hand speed follows a typical profile (bold apricot lines). According to the body mass underestimation hypothesis (left panel), initial underactuation produces a lower, earlier peak speed (light red), necessitating a corrective submovement (light blue) that results in an asymmetric, prolonged profile (purple). The conservative strategy hypothesis (right panel) predicts symmetric slowing with delayed peak speed (red). B) Two-joint arm model and its simulated effective mass. The spatial distribution of effective mass is shown by black curves (solid: normal gravity; dashed: 30% mass underestimation in microgravity). Colored lines indicate effective mass values for our three experimental target directions (45°, 90°, 135°). C & F) Representative speed and acceleration profiles simulated using effective masses from B. Profiles shown for normal (solid) and underestimated (dashed) body mass in the 90° direction demonstrate reduced, earlier peaks with mass underestimation. D-E & G-H) Simulated kinematic parameters across movement directions under normal (solid) and underestimated (dashed) mass conditions. Mass underestimation consistently leads to reduced amplitude and earlier timing of speed/acceleration peaks across all directions. Model simulation details are provided in the Methods.

Revised version: A) Speed profile predictions under two hypotheses. Under normal gravity, hand speed follows a typical profile (bold apricot lines). According to the body mass underestimation hypothesis (left panel), initial underactuation produces a lower, earlier peak speed (light red), necessitating a corrective submovement (light blue) that results in an asymmetric, prolonged profile (purple). The conservative strategy hypothesis (right panel) predicts symmetric slowing with delayed peak speed (red). B) Two-joint arm model and its simulated effective mass. The spatial distribution of effective mass is shown by black curves (solid: normal gravity; dashed: 30% mass underestimation in microgravity). Colored lines indicate effective mass values for our three experimental target directions (45°, 90°, 135°). C & F) Representative speed and acceleration profiles simulated using effective masses from B. Profiles shown for normal (solid) and underestimated (dashed) body mass in the 90° direction demonstrate reduced, earlier peaks with mass underestimation. D-E & G-H) Simulated kinematic parameters across movement directions under normal (solid) and underestimated (dashed) mass conditions. Mass underestimation consistently leads to reduced amplitude and earlier timing of speed/acceleration peaks across all directions. Model simulation details are provided in the Methods.

Earlier version: To resolve the theoretical debate, we leveraged a unique biomechanical property of arm movements: the natural variation in effective limb mass across movement directions, known as anisotropic inertia (Mussa-Ivaldi et al., 1985). Based on a two-joint arm model (Li & Todorov, 2004), reaching movements toward targets in the frontal-medial direction involve greater effective mass compared to those in the frontal-lateral direction. Based on the motor utility theory, both the magnitude and timing of peak speed and acceleration should be systematically modulated by target direction (Shadmehr et al., 2016). Similarly, if astronauts underestimate their limb mass by approximately 30% in microgravity (Crevecoeur et al., 2014), we would expect reduced peak speed/acceleration and earlier occurrence of these peaks across all the directions. Here we examined twelve taikonauts before, during, and after their 3- or 6-month missions on the China Space Station (CSS) and compared their performance to age-matched ground controls. Participants performed rapid reaching movements to targets specifically chosen to engage different effective masses. Our approach allows us to dissociate the effects of mass underestimation from strategic control by examining how movement kinematics vary with direction-dependent effective mass.

Revised version: To help distinguish between these theories, we leveraged a unique biomechanical property of arm movements: the natural variation in effective limb mass across movement directions, known as anisotropic inertia (Mussa-Ivaldi et al., 1985). Based on a two-joint arm model (Li & Todorov, 2004), reaching movements toward targets in the frontal-medial direction involve greater effective mass compared to those in the frontal-lateral direction. Based on the motor utility theory, both the magnitude and timing of peak speed and acceleration should be systematically modulated by target direction (Shadmehr et al., 2016). Similarly, if astronauts underestimate their limb mass by approximately 30% in microgravity (Crevecoeur et al., 2014), we would expect reduced peak speed/acceleration and earlier occurrence of these peaks across all the directions. Here we examined twelve taikonauts before, during, and after their 3- or 6-month missions on the China Space Station (CSS) and compared their performance to age-matched ground controls. Participants performed rapid reaching movements to targets specifically chosen to engage different effective masses. Our approach allows us to dissociate the effects of mass underestimation from strategic control by examining how movement kinematics vary with direction-dependent effective mass.

Results

We compared the performance of 12 taikonauts to that of age-matched ground controls across multiple test sessions (Figure 6B and Table S1). Participants performed rapid reaching movements toward one of three targets on a tablet, presented in pseudorandom order to minimize movement automation and ensure active planning for each reach. Participants were instructed to reach the target with both speed and accuracy, pausing briefly before initiating the return. The emphasized requirement of moving fast would critically test whether movement slowing is ubiquitous in microgravity. The three target directions (45°, 90°, and 135°) were chosen to vary the effective mass of the moving limb, as verified through biomechanical simulations based on the two-joint arm model ((Li & Todorov, 2004); Figure 1B). Our model simulations not only demonstrated the direction-dependent variation in limb mass but also predicted how actual and misperceived limb mass would affect reaching kinematics (Figure 1C-H). These predictions were derived by combining the movement utility framework for estimating planned movement duration ((Shadmehr et al., 2016); Figure S1A-B) with a forward optimal controller for simulating reaching trajectories ((Todorov, 2005); Figure 1-figure supplement 1C-D). Notably, the model predicted that underestimating limb mass would lead to reduced and earlier peak speed/acceleration across directions (Figure 1C-H).

Movement Accuracy and Reaction Time Remain Intact in Microgravity

All participants successfully completed the task. Invalid trials, comprising early initiation (RT < 150 ms), late initiation (RT > 400 ms), or measurement failures, accounted for only 1.87% of trials in the experimental group and 0.70% in the control group. The percentage of invalid trials remained stable across the pre-in and post-flight phases (Friedman’s tests, all p > 0.1).

Reaching accuracy remained high despite the stringent speed requirement of completing movements within 650 ms after target appearance (Figure 2-figure supplement 2). Analysis of endpoint error using a 3 (target direction) × 3 (phase) repeated-measures ANOVA revealed no significant differences across phases for either group. While the experimental group showed slightly better accuracy for the 45° target, this difference was minimal (0.04 cm over a 12 cm reach distance).

Removed in the revised version: To introduce sufficient variability in reaction times for assessing speed–accuracy trade-offs (Sutter et al., 2021), a beep sound was played upon the target appearance in half of the trials. The random auditory signal tended to accelerate responses, thereby widening the distribution of reaction times in our dataset. The beep effectively reduced median reaction times by approximately 38 ms (242 ± 49 ms vs. 281 ± 50 ms; Figure 2A; see control group data in Figure 2-figure supplement 1A). To evaluate microgravity effects, we pooled data from both beep conditions and conducted a 3 (direction) × 3 (phase) repeated-measures ANOVA. The experimental group showed significant main effects of direction (F(2,22) = 25.260, p < 0.001, partial η2 = 0.697) and phase (F(2,22) = 5.820, p = 0.014, partial η2 = 0.346), without interaction (F(4,44) = 2.246, p = 0.111). Post-hoc analyses revealed slower reaction times for the 45° direction compared to both 90° (p < 0.001, d = 0.293) and 135° (p = 0.003, d = 0.284). Contrary to the conservative-strategy hypothesis, taikonauts did not show generalized slowing in reaction time; they actually had faster reaction times during spaceflight, incompatible with a generalized slowing strategy. Reactions were faster during the in-flight phase compared to pre-flight (p = 0.037, d = 0.333), with no significant difference between in-flight and post-flight phases (p = 0.127). Similar effects were observed in the control group (Figure 2-figure supplement1A), suggesting that the reduced reaction times likely reflect practice effects rather than microgravity influence. Thus, participants maintained, and even improved, their ability to rapidly initiate movements in microgravity with reaching accuracy.

Reaction time and movement duration in the experimental group.

A) Median reaction times plotted by experimental phase (x-axis) and target direction (different colors). Dashed lines and dot lines represent the “beep” and “no beep” conditions. B) Average movement durations shown in the same format. Data are combined across beep conditions due to no significant differences between them. In both panels, error bars denote standard error across participants. Asterisks indicate significance levels for post-hoc phase comparisons (*p < 0.05, **p < 0.01, ***p < 0.001).

Added in the revised version: To introduce sufficient variability in reaction times for assessing speed–accuracy trade-offs (Sutter et al., 2021), a beep sound was played upon the target appearance in half of the trials. The random auditory signal tended to accelerate responses, thereby widening the distribution of reaction times in our dataset. The beep effectively reduced median reaction times by approximately 38 ms (242 ± 49 ms vs. 281 ± 50 ms; Figure 2A; see control group data in Figure 2-figure supplement 1A). To evaluate microgravity effects, we pooled data from both beep conditions and conducted a 3 (direction) × 3 (phase) repeated-measures ANOVA. The experimental group showed significant main effects of direction (F(2,22) = 25.260, p < 0.001, partial η2 = 0.697) and phase (F(2,22) = 5.820, p = 0.014, partial η2 = 0.346), without interaction (F(4,44) = 2.246, p = 0.111). Post-hoc analyses revealed slower reaction times for the 45° direction compared to both 90° (p < 0.001, d = 1.473) and 135° (p = 0.003, d = 1.427). Contrary to the conservative-strategy hypothesis, taikonauts did not show generalized slowing in reaction time; they actually had faster reaction times during spaceflight, incompatible with a generalized slowing strategy. Reactions were faster during the in-flight phase compared to pre-flight (p = 0.037, d = 0.803), with no significant difference between in-flight and post-flight phases (p = 0.127). Similar effects were observed in the control group (Figure 2-figure supplement 1A), suggesting that the reduced reaction times likely reflect practice effects rather than microgravity influence. Thus, participants maintained, and even improved, their ability to rapidly initiate movements in microgravity with reaching accuracy.

Microgravity leads to slower movement

Earlier version: Although showing no decrement in movement accuracy or RT, the taikonauts exhibited prolonged movement durations (MD) during the in-flight phase compared to both pre- and post-flight phases (Figure 2B). A two-way repeated-measures ANOVA reveals significant main effects of direction (F(2,22) = 62.555, p < 0.001, partial η2 = 0.850) and phase (F(2,22) = 6.761, p = 0.015, partial η2 = 0.381), with no significant interaction (F(4,44) = 1.413, p = 0.263). Post-hoc comparisons showed that MD was significantly prolonged during the in-flight phase (vs. pre-flight: p = 0.014, d = 0.333; vs. post-flight: p = 0.013, d = 0.181). MD was also significantly greater for the 90° and 135° directions—those associated with larger effective mass—compared to the 45° direction (all p < 0.001 in pairwise comparisons). In contrast, the control group showed a similar directional effect but, importantly, did not show a phase effect (Figure 2—figure supplement 1), suggesting that movement slowing is specific to spaceflight.

Revised version: Although showing no decrement in movement accuracy or RT, the taikonauts exhibited prolonged movement durations (MD) during the in-flight phase compared to both pre- and post-flight phases (Figure 2B). A two-way repeated-measures ANOVA reveals significant main effects of direction (F(2,22) = 62.555, p < 0.001, partial η2 = 0.850) and phase (F(2,22) = 6.761, p = 0.015, partial η2 = 0.381), with no significant interaction (F(4,44) = 1.413, p = 0.263). Post-hoc comparisons showed that MD was significantly prolonged during the in-flight phase (vs. pre-flight: p = 0.014, d = 0.866; vs. post-flight: p = 0.013, d = 0.472). MD was also significantly greater for the 90° and 135° directions—those associated with larger effective mass—compared to the 45° direction (all p < 0.001 in pairwise comparisons). In contrast, the control group showed a similar directional effect but, importantly, did not show a phase effect (Figure 2—figure supplement 1), suggesting that movement slowing is specific to spaceflight.

Microgravity leads to reduced peak acceleration/speed and advanced timing of these peaks

The initial movement exhibited signs of underactuation during spaceflight, consistent with the body mass underestimation hypothesis (Figure 2—figure supplement 6). Peak acceleration and peak speed, established measures of feedforward control (Elliott et al., 2010, 2017), both showed significant decreases during the in-flight phase (Figure 3). For peak acceleration, there were significant main effects of phase (F(2,22) = 6.401, p = 0.015, partial η2 = 0.368) and direction (F(2,22) = 75.516, p < 0.001, partial η2 = 0.873), along with an interaction effect (F(4,44) = 4.548, p = 0.016, partial η2 = 0.293). Consistent with model simulations, the 45° direction, associated with a lower effective mass, showed significantly higher peak acceleration than the other two directions (Figure 3A; all p < 0.001 in pairwise comparisons). Peak acceleration decreased significantly during spaceflight across all three directions (simple main effects: 45°, p = 0.020, η2 = 0.543; 90°, p = 0.004, η2 = 0.673; 135°, p = 0.005, η2 = 0.657). Planned contrasts further indicated lower peak acceleration in the in-flight phase compared to pre-flight (45°, p = 0.006, η2 = 0.516; 90°, p = 0.008, η2 = 0.485; 135°, p = 0.003, η2 = 0.560), with marginal recovery in the post-flight phase for two directions (45°, p = 0.071, η2 = 0.267; 90°, p = 0.062, η2 = 0.282) and significant recovery for one direction (135°, p = 0.029, η2 = 0.363). In contrast, the control group displayed only the direction effect (Figure 3—figure supplement 1A-B), confirming that the underactuation was specific to microgravity.

Magnitude changes of peak acceleration and speed during spaceflight.

A) Peak acceleration across experimental phases and movement directions, showing systematic decreases during the in-flight phase. B) Peak speed data presented in the same format, demonstrating parallel changes to peak acceleration. Error bars indicate standard error across participants. Statistical significance levels are shown for both overall phase comparisons (black asterisks) and planned contrasts within each direction (colored asterisks): *p < 0.05, **p < 0.01, ***p < 0.001.

Earlier version: Peak speed showed parallel changes during spaceflight (Figure 3B). Analysis revealed significant main effects of phase (F(2,22) = 7.043, p = 0.009, partial η2 = 0.390) and direction (F(2,22) = 114.420, p < 0.001, partial η2 = 0.912), without interaction (F(4,44) = 1.929, p = 0.161). Peak speed was significantly reduced during the in-flight phase compared to both pre-flight (p = 0.004, d = 0.331) and post-flight phases (p = 0.046, d = 0.199), with evidence of recovery post-flight (pre-vs. post-flight: p = 0.483, d = 0.132). Target direction systematically modulated peak speed (all pairwise comparisons, p < 0.001), with the 45° direction showing highest speeds and 135° lowest—matching predictions based on the movement utility theory. The control group showed only direction effects without phase-related changes (Figure 3—figure supplement 1B), confirming the specificity of these changes to microgravity exposure.

Revised version: Peak speed showed parallel changes during spaceflight (Figure 3B). Analysis revealed significant main effects of phase (F(2,22) = 7.043, p = 0.009, partial η2 = 0.390) and direction (F(2,22) = 114.420, p < 0.001, partial η2 = 0.912), without interaction (F(4,44) = 1.929, p = 0.161). Peak speed was significantly reduced during the in-flight phase compared to both pre-flight (p = 0.004, d = 0.879) and post-flight phases (p = 0.046, d = 0.529), with evidence of recovery post-flight (pre-vs. post-flight: p = 0.483, d = 0.349). Target direction systematically modulated peak speed (all pairwise comparisons, p < 0.001), with the 45° direction showing highest speeds and 135° lowest—matching predictions based on the movement utility theory. The control group showed only direction effects without phase-related changes (Figure 3—figure supplement 1B), confirming the specificity of these changes to microgravity exposure.

Earlier version: The critical test for the competing hypotheses of movement slowing lies in whether peak acceleration and peak speed occur earlier or later in microgravity. The body mass underestimation hypothesis predicts earlier timing, whereas the conservative control strategy hypothesis predicts delayed timing. We observed an overall timing advance of these two measures. For peak acceleration time in the experimental group (Figure 4A), two-way repeated-measures ANOVAs revealed significant main effects of phase (F(2,22) = 5.189, p = 0.024, partial η2 = 0.321) and direction (F(2,22) = 30.375, p < 0.001, partial η2 = 0.734), along with an interaction effect (F(4,44) = 4.503, p = 0.013, partial η2 = 0.290). In line with predictions, peak acceleration appeared significantly earlier in the 45° direction than other directions (45° vs. 90°, p < 0.001, d = 0.304; 45° vs. 135°, p < 0.001, d = 0.271). The phase effect emerged only for directions with higher effective mass (simple main effects, 90°: F(2,10) = 6.660, p = 0.015, η2 = 0.571; 135°: F(2,10) = 8.411, p = 0.007, η2 = 0.627). Planned contrasts confirmed earlier peak acceleration during spaceflight compared to post-flight (90°, p = 0.005, η2 = 0.530; 135°, p = 0.004, η2 = 0.537), with marginal differences versus pre-flight (135°, p = 0.065). The control group showed only directional effects without phase effect (Figure 4—figure supplement 1A).

Revised version: The critical test between the two hypotheses of movement slowing lies in whether peak acceleration and peak speed occur earlier or later in microgravity. The body mass underestimation hypothesis predicts earlier timing, whereas the conservative control strategy hypothesis predicts delayed timing. We observed an overall timing advance of these two measures. For peak acceleration time in the experimental group (Figure 4A), two-way repeated-measures ANOVAs revealed significant main effects of phase (F(2,22) = 5.189, p = 0.024, partial η2 = 0.321) and direction (F(2,22) = 30.375, p < 0.001, partial η2 = 0.734), along with an interaction effect (F(4,44) = 4.503, p = 0.013, partial η2 = 0.290). In line with predictions, peak acceleration appeared significantly earlier in the 45° direction than other directions (45° vs. 90°, p < 0.001, d = 1.675; 45° vs. 135°, p < 0.001, d = 1.491). The phase effect emerged only for directions with higher effective mass (simple main effects, 90°: F(2,10) = 6.660, p = 0.015, η2 = 0.571; 135°: F(2,10) = 8.411, p = 0.007, η2 = 0.627). Planned contrasts confirmed earlier peak acceleration during spaceflight compared to post-flight (90°, p = 0.005, η2 = 0.530; 135°, p = 0.004, η2 = 0.537), with marginal differences versus pre-flight (135°, p = 0.065). The control group showed only directional effects without phase effect (Figure 4—figure supplement 1A).

Temporal analysis of peak acceleration and peak speed during spaceflight.

A-B) Time to peak acceleration and peak speed across experimental phases and movement directions. C-D) Relative timing analysis showing these same peaks as a percentage of total movement duration, revealing more pronounced microgravity effects than absolute timing measures. Error bars indicate standard error across participants. Asterisks denote significance levels for planned phase comparisons (*p < 0.05, **p < 0.01, ***p < 0.001). Note the systematic modulation of timing by phase, particularly for high effective mass conditions.

Given that overall movement duration is prolonged in microgravity, a more direct way to assess the symmetry of the speed profile is to compute the relative timing of speed/acceleration peaks as a percentage of MD. Indeed, the relative timing of peak acceleration showed the same reduction in microgravity with a more pronounced effect than absolute timing (Figure 4C). A two-way repeated-measures ANOVA revealed significant main effects of phase (F(2,22) = 15.414, p < 0.001, partial η2 = 0.584) and interaction (F(2,22) = 7.223, p = 0.002, partial η2 = 0.396), while the main effect of direction was marginally significant (F(2,22) = 3.374, p = 0.053, partial η2 = 0.235). Phase effects were significant for directions with higher effective mass (simple main effects: 90°, F(2,10) = 13.327, p = 0.002, η2 = 0.727; 135°, F(2,10) = 24.574, p < 0.001, η2 = 0.831) but not for 45° (F(2,10) = 0.609, p = 0.563). Planned contrasts confirmed significantly earlier timing during spaceflight compared to both pre- and post-flight phases (all p < 0.002 for 90° and 135° directions). None of these effects appeared in the control group (Figure 4—figure supplement 1C).

Peak speed also showed a timing advance, but only for relative timing. For absolute timing (Figure 4B), a two-way repeated-measures ANOVA revealed a significant main effect of direction (F(2,22) = 57.428, p < 0.001, partial η2 = 0.839) and a marginal interaction (F(4,44) = 2.233, p = 0.095, partial η2 = 0.169), without a phase effect (F(2,22) = 0.093, p = 0.864). In contrast, relative timing analysis (Figure 4D) showed significant effects of phase (F(2,22) = 8.941, p = 0.004, partial η2 = 0.448), direction (F(2,22) = 4.991, p = 0.030, partial η2 = 0.312), and their interaction (F(4,44) = 4.755, p = 0.015, partial η2 = 0.302). The phase effect appeared only for directions with higher effective mass (90°: F(2,10) = 17.278, p = 0.001, η2 = 0.776; 135°: F(2,10) = 6.189, p = 0.018, η2 = 0.553). Planned contrasts revealed significantly earlier relative timing during spaceflight versus pre-flight for both 90° (p < 0.001, η2 = 0.770) and 135° (p = 0.007, η2 = 0.496) directions, with significant post-flight recovery in 90° (p = 0.003, η2 = 0.575). The control group showed only directional effects (Figure 4—figure supplement 1B & 1D). Thus, after accounting for the prolongation of movement duration, we observed a temporal advance of peak acceleration and speed during spaceflight.

Movement slowing is associated with corrective submovements

The feedforward component of movement, captured by our model simulations, accounts for only the initial phase of reaching. With peak acceleration and speed occurring at approximately 80 ms and 150 ms after movement onset, and total movement duration spanning 300-400 ms, participants had ample opportunity for feedback-based corrections (Dimitriou et al., 2013). These corrections, if present, would manifest as secondary submovements (D. E. Meyer et al., 1988). While optimal reaching movements typically show symmetrical speed profiles (Flash & Hogan, 1985), the presence of corrective submovements creates characteristic right-skewed speed profiles (Elliott et al., 2010; D. E. Meyer et al., 1988; Woodworth, 1899). To identify such corrections, we employed an established decomposition method (Rohrer & Hogan, 2006) to detect trials containing both primary and secondary submovements (two-peak trials; Figure 5A).

Submovements analysis revealed changes in corrective movements during spaceflight.

A) Decomposition of a representative two-peak speed profile, illustrating primary and corrective submovements separated by the inter-peak interval (IPI). B) Proportion of movements showing corrective submovements across phases and directions. C) Magnitude of feedback corrections quantified by IPI, showing direction-dependent increases in microgravity. D) Linear relationship between feedback correction changes (ΔIPI) and movement slowing (ΔMD). Each participant (shown in different colors) contributed multiple data points from different phases and directions. E-F) Primary submovement characteristics: peak speed amplitude and timing. Error bars denote standard error across participants; asterisks indicate significance (*p < 0.05, **p < 0.01, ***p < 0.001).

Analysis revealed a striking increase in corrective movements during spaceflight. The percentage of trials showing two distinct submovements increased significantly during the in-flight phase (Figure 5B), evidenced by a main effect of phase (F(2,22) = 13.969, p < 0.001, partial η2 = 0.559) and a marginal phase-by-direction interaction (F(2,22) = 2.642, p = 0.054, partial η2 = 0.194). This effect was specific to microgravity exposure, as the control group showed no significant changes across phases (all p > 0.05; Figure 5—figure supplement 1A). Notably, the increase in corrective movements was most pronounced for directions with higher effective mass (simple main effects: 90°, F(2,10) = 11.426, p = 0.003, η2 = 0.696; 135°, F(2,10) = 7.292, p = 0.011, η2 = 0.593).

To understand how feedback corrections contribute to movement slowing, we analyzed the inter-peak interval (IPI) between primary and corrective submovements, an established measure of feedback-based control (Craik, 1947; D. Meyer et al., 2018; D. E. Meyer et al., 1988). The taikonauts showed markedly longer IPIs during spaceflight (Figure 5C), with a two-way ANOVA revealing significant effects of phase (F(2,22) = 11.150, p = 0.001, partial η2 = 0.503), direction (F(2,22) = 4.721, p = 0.026, partial η2 = 0.300), and their interaction (F(2,22) = 6.559, p = 0.002, partial η2 = 0.374). Critically, IPI increases were confined to directions with greater effective mass (90°, F(2,10) = 11.573, p = 0.002, η2 = 0.698; 135°, F(2,10) = 13.371, p = 0.001, η2 = 0.728). The control group exhibited only directional effects (F(2,22) = 19.277, p < 0.001, partial η2 = 0.637; Figure 5—figure supplement 1B) without phase-related changes (all p > 0.13).

To establish whether these enlarged corrective movements explain the overall movement slowing, we examined how changes in IPI (ΔIPI) predict changes in movement duration (ΔMD) during spaceflight (Figure 5D). A Linear Mixed Model analysis incorporating ΔIPI, phase transition (pre-to-in and post-to-in), and direction revealed that ΔIPI significantly predicted ΔMD (β = 0.458 [0.144, 0.771], t(67) = 2.915, p = 0.005, partial-R2 = 0.103). Phase transition also showed significance (β = −0.01 [−0.017, −0.001], t(67) = −2.696, p = 0.009, partial-R2 = 0.086), with stronger effects for pre-to-in versus post-to-in transitions, indicating incomplete recovery post-flight. Movement direction showed no significant effect (all p > 0.42), suggesting that the relationship between corrective movements and duration slowing generalizes across reaching directions.

The analysis of primary submovements revealed even stronger evidence for feedforward control changes in microgravity. By focusing on just the primary submovement in two-peak trials, we found more pronounced effects than when analyzing the overall movement (comparing Figure 5E to Figure 5—figure supplement 1C). The primary submovement showed a robust reduction in peak speed during spaceflight (main effect of phase: F(2,22) = 10.363, p = 0.001, partial η2 = 0.485) that was consistent across directions (interaction: F(4,44) = 0.382, p = 0.711; Figure 5E). More tellingly, while the overall movement showed timing changes only in relative measures, the primary submovement exhibited earlier peak speed timing in absolute terms (phase-by-direction interaction: F(2,22) = 4.264, p = 0.012, partial η2 = 0.279). This temporal advance was most evident in directions with higher effective mass (planned contrasts vs. post-flight: 135°: p = 0.008, η2 = 0.615; 90°: p = 0.076, η2 = 0.403) but absent in the 45° direction (p = 0.188, η2 = 0.284). The selective nature of the primary submovement in both magnitude and timing (comparing Figure 4B and Figure 5F) provides compelling evidence for mass-dependent changes in feedforward control. The control group showed only the expected directional variations (F(2,22) = 46.259, p < 0.001, partial η2 = 0.808; Figure 5—figure supplement 1D) without any phase effects (all p > 0.13), confirming these changes as specific to microgravity exposure.

The microgravity effects on movement kinematics are directionally dependent

Reaching movements are inherently directionally dependent due to anisotropies in limb dynamics (Gordon et al., 1994; Shadmehr et al., 2016). Even before microgravity exposure, key kinematic variables—peak speed, peak acceleration, and their timing—exhibited consistent directional differences, as shown in Figure 1C–H. This baseline direction-dependency suggests that effective mass varies across movement directions and influences movement execution under normal gravity. Building on this observation, we tested whether microgravity-induced changes would also follow a direction-specific pattern, as predicted by the mass underestimation hypothesis.

To assess microgravity-induced changes, we quantified the directional differences (Δ) in peak speed, peak acceleration, and their timing, comparing in-flight measurements to pre-flight baselines. The comparison between in-flight and post-flight was not included in this analysis since post-flight recovery was not complete (see peak speed and peak acceleration changes in Figure 3). As shown in Figure 6, these directional effects were evident across all kinematic metrics. For both Δ peak speed and Δ peak acceleration, the magnitude of change followed a consistent ranking across directions, with the largest effects at 45°, intermediate at 90°, and smallest at 135° (i.e., 45° > 90° > 135°). Similarly, for timing metrics—Δ peak speed time and Δ peak acceleration time—the shortest delays occurred at 45°, while 90° and 135° showed comparably larger shifts. Note the 45° did not show significant timing advance, as previously shown in Figure 5F.

Direction-dependent effects of microgravity on peak kinematics and their timing.

Added in the revised version: Left: model outputs. Right: experimental data shown as Δ relative to the in-flight session; solid bars = Δ(in − pre) and semi-transparent bars = Δ(in − post). Colors encode direction consistently across panels (e.g., 45° = darker hue, 90° = medium, 135° = lighter/orange). Panels (clockwise from top-left): Δ peak speed (cm/s), Δ peak speed time (ms), Δ peak acceleration time (ms), and Δ peak acceleration (cm/s2). Bars are group means; error bars denote standard error across participants.

Removed in the revised version: Each panel shows the change (Δ) relative to the in-flight session. Colors represent movement directions consistently across panels. Panels A–D show, respectively: (A) change in peak speed (cm/s), (B) change in peak speed time (ms), (C) change in peak acceleration time (ms), and (D) change in peak acceleration (cm/s2). Bars represent group means; error bars denote standard errors across participants.

Earlier version: Overall, the direction-dependent changes were roughly consistent with model predictions based on the mass underestimation hypothesis (Figure 1 D, E, G and H). The model simulation similarly indicates that amplitude changes (Δ peak speed, Δ peak acceleration) and timing advances are rank ordered by movement directions. To quantify the correspondence between model predictions and experimental measurements, we performed repeated-measures correlation analyses between simulated and observed Δ values across directions (Bakdash & Marusich, 2017). Significant within-subject correlations were found for Δ peak speed time (rrm = 0.627, p < 0.001), Δ peak acceleration time (rrm = 0.591, p = 0.002), and Δ peak acceleration (rrm = 0.573, p = 0.003), while Δ peak speed did not reach significance (rrm = 0.334, p = 0.103). Together, these results demonstrate that the microgravity effect is direction-specific and the rank ordering of directions is broadly aligned with the predictions of the mass underestimation model.

Revised version: Overall, the direction-dependent changes were roughly consistent with model predictions based on the mass underestimation hypothesis (Figure 6, comparing left and right panels). The model simulation similarly indicates that amplitude changes (Δ peak speed, Δ peak acceleration) and timing advances are rank ordered by movement directions. While the absolute magnitudes differed between model and data, the directional trends were broadly aligned. To assess the overall consistency between model predictions and experimental measurements, we performed repeated-measures correlation analyses between simulated and observed Δ values across directions (Bakdash & Marusich, 2017). Significant within-subject correlations were found for Δ peak speed time (rm = 0.627, p < 0.001), Δ peak acceleration time (rm = 0.591, p = 0.002), and Δ peak acceleration (rm = 0.573, p = 0.003), while Δ peak speed did not reach significance (rm = 0.334, p = 0.103). Together, these results suggest that the microgravity effect is direction-specific and the rank ordering of directions is broadly aligned with the predictions of the mass underestimation model. We note that these correlations evaluate the directional trend rather than the absolute magnitude of the effects; a precise quantitative match is not expected given the simplifications of the two-joint arm model.

Speed-accuracy trade-offs remain intact during spaceflight

The speed-accuracy trade-off, a fundamental indicator of motor control efficiency, remained largely intact during spaceflight. We evaluated this trade-off using two complementary approaches. First, following the traditional Fitts’ framework =, we examined the relationship between movement duration and endpoint dispersion, which was not negatively affected by microgravity (Figure 2—figure supplement 3). Second, we analyzed the trade-off between reaction time and initial movement variance (Figure 2—figure supplement 4), which reflects the quality of movement planning (Sutter et al., 2021). This relationship persisted across all phases in both groups. Thus, while microgravity altered movement execution through mass underestimation, it did not fundamentally disrupt the sensorimotor system’s capacity to regulate the speed-accuracy trade-off. Additionally, in both groups, we found a significant negative correlation between movement duration (MD) and reaction time (RT), both across and within individuals (Figure 2—figure supplement 5). This finding indicates that participants moved faster when their RT was slower, and vice versa. This flexible motor adjustment, likely in response to the task requirement for rapid movements, remained consistent during spaceflight.

Discussion

Earlier version: Long-term exposure to microgravity presents a unique environment that cannot be replicated on Earth, offering invaluable insights into human sensorimotor adaptation. Understanding how microgravity affects motor performance is crucial both for ensuring successful space exploration and for advancing our fundamental knowledge of motor control principles. Our study elucidates the mechanisms underlying one of the most distinctive motor signatures in microgravity: the slowing of goal-directed actions during spaceflight. While taikonauts maintained reaching accuracy and rapid reaction times, their movements exhibited systematic changes in microgravity characterized by reduced peak speed and acceleration, along with advanced timing of these peaks. These kinematic alterations persisted across repetitive trials within a session and also persisted across multiple sessions during spaceflight (Figure 4—figure supplement 2), contradicting the hypothesis that astronauts adopt a conservative control strategy to maintain stability. Instead, they align precisely with our hypothesis that individuals underestimate limb effective mass, leading to underactuated initial movements.

Revised version: Long-term exposure to microgravity presents a unique environment that cannot be replicated on Earth, offering invaluable insights into human sensorimotor adaptation. Understanding how microgravity affects motor performance is crucial both for ensuring successful space exploration and for advancing our fundamental knowledge of motor control principles. Our study investigates the mechanisms underlying one of the most distinctive motor signatures in microgravity: the slowing of goal-directed actions during spaceflight. While taikonauts maintained reaching accuracy and rapid reaction times, their movements exhibited systematic changes in microgravity characterized by reduced peak speed and acceleration, together with direction-dependent advances in peak timing, particularly for movements involving higher effective mass. These persistent kinematic alterations (Figure 4—figure supplement 2) are not readily explained by a simple conservative-control account alone. Instead, they are consistent with an underactuation pattern that could arise, at least in part, from an underestimation of limb effective mass in microgravity.

Earlier version: These changes followed model predictions across movement directions with varying effective masses due to limb biomechanics, confirming that the representation of body mass indeed affects movement control during spaceflight. Beyond the initial underactuation caused by altered feedforward control, we observed increased feedback-based corrective submovements that predicted movement duration prolongation in spaceflight. These effects were specific to microgravity exposure, occurring only during spaceflight with marked recovery post-flight, while ground controls showed no comparable changes. Notably, the taikonauts’ fundamental motor control capabilities remained largely intact, as evidenced by preserved movement accuracy, reaction time, and speed-accuracy trade-offs.

Revised version: These changes were broadly consistent with model predictions across movement directions with varying effective masses due to limb biomechanics, suggesting that the representation of body mass may affect movement control during spaceflight. Beyond the initial underactuation associated with altered feedforward control, we observed increased feedback-based corrective submovements that predicted movement duration prolongation in spaceflight. These effects were specific to microgravity exposure, occurring only during spaceflight with marked recovery post-flight, while ground controls showed no comparable changes. Notably, the taikonauts’ fundamental motor control capabilities remained largely intact, as evidenced by preserved movement accuracy, reaction time, and speed-accuracy trade-offs.

The underactuation observed in initial movements under microgravity extends findings from parabolic flight studies (Crevecoeur et al., 2014; Papaxanthis et al., 2005). When individuals performed reaching movements with handheld weights during parabolic flight-induced gravity alterations, their reaching trajectories and grip forces showed systematic changes suggesting mass misperception (Bock, 1998; Bock et al., 1996a, 1996b; Crevecoeur et al., 2014; Papaxanthis et al., 2005). In zero-gravity conditions, goal-directed arm movements were slower (Papaxanthis et al., 2005), and grip-load force coordination was altered (Crevecoeur et al., 2010, 2014), consistent with mass underestimation. Conversely, hypergravity (1.8g) induced increased peak acceleration and speed in reaching movements, along with elevated grip force during object manipulation, suggesting mass overestimation (Bock, 1998; Bock et al., 1992, 1996a, 1996b). Notably, modified grip-load force coordination reflects changes in feedforward control (Flanagan & Wing, 1997; Johansson & Edin, 1993), indicating that mass misestimation influences action planning. However, while parabolic flight studies capture only transient adaptations to rapid gravity transitions, our study targets general slowing during spaceflight and reveals the enduring effects of sustained microgravity exposure.

Earlier version: We also considered alternative explanations beyond mass underestimation. First, we tested whether a change in the cost-function of our optimal control model, which simulates a more conservative control, could account for movement slowing. This was implemented as a uniform rescaling of the state and control weights (Q and R) in the cost function (Crevecoeur et al., 2010). Increasing the scaling factors (α) will simulate the increase in control effort penalty and a more conservative control. However, while this manipulation reduced peak velocity and acceleration, it also delayed their occurrence, a pattern opposite to the observed effect. Thus, changing the cost function alone is unable to explain away the movement slowing. Details of the cost-function simulations are provided in the Supplementary Materials (Supplementary Notes 1 and Figure 1—figure supplements 2 and 3). Second, muscle weakness, a common effect during prolonged spaceflight, might also contribute to movement slowing. However, recent large-sample studies suggest that muscle strength reduction is mainly in lower-limb muscles that counter gravitational pull, not in upper-body muscles (Scott et al., 2023). Furthermore, the torques required in our task (∼2 N·m for ∼12 cm reaches) are relatively small, it is unlikely that muscle weakness alone accounts for the observed kinematics. Third, finger-screen friction might slow down the reaching movements. However, during typical interaction with a touch screen, frictional forces are modest, generally ranging from 0.1 to 0.5 N (Ayyildiz et al., 2018), and directional variations are likely even smaller. These values are minimal compared to the 10–15 N required to accelerate the arm during reaching, making frictional anisotropy an unlikely source of the directional kinematic changes. Fourth, body stabilization in microgravity might also contribute to movement slowing in space. In fact, our participants were securely strapped at the feet and used the non-reaching hand to grasp a fixed bar for support. Importantly, the primary effects—reduced peak acceleration and earlier timing—emerged within the first ∼90–100 ms of movement onset, well before any instability-related corrections could occur. Thus, insufficient stabilization is unlikely to account for the observed movement changes. Finally, Coriolis and centripetal torques, omitted from our simplified model, might also contribute to movement slowing or direction effect we observed. Previous work has suggested that such effects may be non-negligible in fast planar reaching (Hollerbach & Flash, 1982). To examine their potential contribution, we incorporated Coriolis and centripetal torques terms in our extended model (see Supplementary Notes 2 and Figure 1—figure supplement 4 and 5 for details). These torques were also directionally anisotropic, but their magnitudes were too small to meaningfully affect movement utility or effort computation in our model. Thus, Coriolis and centripetal torques are unlikely to explain the large and systematical effects we observed.

Revised version: We also considered alternative explanations beyond mass underestimation. First, we tested whether a change in the cost-function of our optimal control model, which simulates a more conservative control, could account for movement slowing. This was implemented as a uniform rescaling of the state and control weights (Q and R) in the cost function (Crevecoeur et al., 2010). Increasing the scaling factors (α) will simulate the increase in control effort penalty and a more conservative control. However, while this manipulation reduced peak velocity and acceleration, it also delayed their occurrence, a pattern opposite to the observed effect. Thus, a uniform change in the cost function alone does not reproduce the observed combination of reduced peak amplitudes and advanced peak timing. Details of the cost-function simulations are provided in the Supplementary Materials (Supplementary Notes 1 and Figure 1—figure supplements 2 and 3). Second, changes in neuromuscular properties may also contribute to the observed underactuation. In microgravity, tonic muscle activity is likely reduced due to the absence of sustained antigravity demands, and descending vestibular inputs to motor neurons are diminished (Fisk et al., 1993). These changes could lower motor neuron excitability and alter muscle activation dynamics, potentially resulting in a weaker and slower initial response to an otherwise equivalent motor command—even in the absence of mass misestimation. While our simplified model treats the actuators as ideal torque generators and therefore does not capture these neuromuscular factors, such factors alone do not straightforwardly explain the direction-dependent pattern observed here. In the absence of additional direction-specific mechanisms, changes in tonic muscle activity, vestibular drive, or reflex gain would be expected to influence movement more generally across directions. The stronger effects observed in directions involving higher effective mass are therefore more naturally captured by a misrepresentation of inertial properties, although neuromuscular changes may also have contributed to the overall underactuation. Third, muscle weakness, a common effect during prolonged spaceflight, might also contribute to movement slowing. However, recent large-sample studies suggest that muscle strength reduction is mainly in lower-limb muscles that counter gravitational pull, not in upper-body muscles (Scott et al., 2023). Furthermore, the torques required in our task (~2 N·m for ~12 cm reaches) are relatively small, making it unlikely that muscle weakness is a primary factor of the observed kinematics, though a minor contribution cannot be excluded. Fourth, finger-screen friction might slow down the reaching movements, particularly in microgravity where astronauts must actively press on the screen to maintain contact in the absence of gravitational loading on the hand. However, during typical interaction with a touch screen, frictional forces are modest, generally ranging from 0.1 to 0.5 N (Ayyildiz et al., 2018), and directional variations are likely even smaller. These values are minimal compared to the 10–15 N required to accelerate the arm during reaching, making frictional anisotropy an unlikely source of the directional kinematic changes. Fifth, body stabilization in microgravity might also contribute to movement slowing in space. In fact, our participants were securely strapped at the feet and used the left hand to grasp a fixed bar for support, providing multi-point stabilization of the body. If trunk displacement were the primary source of the observed effects, it would be expected to uniformly attenuate all kinematic measures, particularly in directions where the reaction force exerts a larger torque on the trunk. However, the 45° direction—where trunk perturbation is expected to be minimal according to this account—showed significant changes in movement duration, peak acceleration, and peak speed, while not showing significant timing changes. This dissociation suggests that trunk instability alone cannot account for the observed pattern. Nonetheless, we did not directly measure trunk or shoulder kinematics during the experiment, and we cannot entirely rule out that small trunk displacements contributed to some extent. Future studies would benefit from recording trunk and shoulder motion to more definitively address this factor. Finally, Coriolis and centripetal torques, omitted from our simplified model, might also contribute to movement slowing or direction effect we observed. Previous work has suggested that such effects may be non-negligible in fast planar reaching (Hollerbach & Flash, 1982). To examine their potential contribution, we incorporated Coriolis and centripetal torques terms in our extended model (see Supplementary Notes 2 and Figure 1—figure supplement 4 and 5 for details). These torques were also directionally anisotropic, but their magnitudes were too small to meaningfully affect movement utility or effort computation in our model. Thus, Coriolis and centripetal torques are unlikely to explain the large and systematical effects we observed.

Earlier version: Our results strongly challenge the conservative control strategy hypothesis, which posits that movement slowing reflects a generalized adaptation to address stability or safety concerns during spaceflight (Bock, 1998). Contrary to this hypothesis, taikonauts demonstrated faster reaction times during spaceflight—incompatible with a generalized slowing strategy. More critically, if the sensorimotor system adopted a conservative strategy with pre-planned longer movement durations, the speed profile should maintain symmetry with delayed peak acceleration and speed. Instead, we observed left-skewed speed profiles with earlier peak occurrence, directly contradicting the predictions of strategic slowing. This temporal change was more pronounced if primary submovement is isolated to pinpoint the feedforward control of the reaching since the overall speed profile is affected by the later feedback-based corrections. Previous studies reported inconsistent findings regarding speed profile symmetry, including advanced (Sangals et al., 1999), delayed (Fowler et al., 2008), or unchanged (Berger et al., 1997; Mechtcheriakov et al., 2002) peak speed timing during spaceflight. These discrepancies likely reflect methodological differences, including task variations and technical limitations. For example, two studies used joystick-based pointing tasks without involving limb movement (Fowler et al., 2008; Sangals et al., 1999), which arguably makes them less suitable for examining mass underestimation effects. Mechtcheriakov and colleagues used an outstretched arm task, but their limited sampling rate of 25 Hz may have hindered precise quantification of peak speed, preventing submovement analysis. Additionally, our larger sample size (n = 12) may have enhanced detection of microgravity effects compared to previous studies with typically 3-7 participants (Berger et al., 1997; Bock et al., 2001, 2003, 2010; Fowler et al., 2008; Mechtcheriakov et al., 2002; Papaxanthis et al., 1998; Weber & Stelzer, 2022) or single-participant designs (Heuer et al., 2003; Manzey et al., 1998; Sangals et al., 1999).

Revised version: Our results challenge a simple uniform conservative-control account, which posits that movement slowing reflects a generalized adaptation to address stability or safety concerns during spaceflight (Bock, 1998). Arguing against a purely strategic account, taikonauts demonstrated faster reaction times during spaceflight—inconsistent with a generalized slowing strategy. Furthermore, if the sensorimotor system adopted a conservative strategy with pre-planned longer movement durations, the speed profile should maintain symmetry with delayed peak acceleration and speed. Instead, we observed left-skewed speed profiles with earlier peak occurrence for the 90° and 135° directions, where effective mass is greater. This pattern is more consistent with mass-dependent under-actuation than with a uniform strategic adjustment, although we note that the 45° direction did not show a significant timing advance, leaving open the possibility that strategic factors may also play a role in some conditions. This temporal change was more pronounced when the primary submovement was isolated to pinpoint the feedforward control of the reaching since the overall speed profile is affected by the later feedback-based corrections. Previous studies reported inconsistent findings regarding speed profile symmetry, including advanced (Sangals et al., 1999), delayed (Fowler et al., 2008), or unchanged (Berger et al., 1997; Mechtcheriakov et al., 2002) peak speed timing during spaceflight. These discrepancies likely reflect methodological differences, including task variations and technical limitations. For example, two studies used joystick-based pointing tasks without involving limb movement (Fowler et al., 2008; Sangals et al., 1999), which arguably makes them less suitable for examining mass underestimation effects. Mechtcheriakov and colleagues used an outstretched arm task, but their limited sampling rate of 25 Hz may have hindered precise quantification of peak speed, preventing submovement analysis. Additionally, our larger sample size (n = 12) may have enhanced detection of microgravity effects compared to previous studies with typically 3-7 participants (Berger et al., 1997; Bock et al., 2001, 2003, 2010; Fowler et al., 2008; Mechtcheriakov et al., 2002; Papaxanthis et al., 1998; Weber & Stelzer, 2022) or single-participant designs (Heuer et al., 2003; Manzey et al., 1998; Sangals et al., 1999).

Earlier version: The underestimation of body mass in microgravity likely arises through two complementary mechanisms: degraded proprioceptive feedback and misinterpretation of weight-related cues. Microgravity impairs proprioception by reducing both muscle spindle sensitivity (Lackner & DiZio, 1992) and joint receptor responsiveness (Proske & Weber, 2023)—inputs that are crucial for weight perception (Proske & Gandevia, 2012). Furthermore, the sensorimotor system may interpret weight unloading as reduced body mass rather than environmental change (Bock et al., 1996b). Recent studies demonstrate that the sensorimotor system probabilistically infers body-versus-environment causality based on available sensory cues (Berniker & Kording, 2008; Fercho & Baugh, 2014; Kluzik et al., 2008; Kong et al., 2017; Wilke et al., 2013). In normal gravity, consistent mass-weight relationships lead to heavy reliance on weight-related cues for mass estimation (Bock, 1998; Bock et al., 1996b; Sangals et al., 1999). In microgravity, however, reduced weight cues become misleading and may systematically bias mass perception, affecting feedforward control of movement.

Revised version: If mass underestimation contributes to the observed underactuation, it likely arises through two complementary mechanisms: degraded proprioceptive feedback and misinterpretation of weight-related cues. Microgravity impairs proprioception by reducing both muscle spindle sensitivity (Lackner & DiZio, 1992) and joint receptor responsiveness (Proske & Weber, 2023)—inputs that are crucial for weight perception (Proske & Gandevia, 2012). Furthermore, the sensorimotor system may interpret weight unloading as reduced body mass rather than environmental change (Bock et al., 1996b). Recent studies demonstrate that the sensorimotor system probabilistically infers body-versus-environment causality based on available sensory cues (Berniker & Kording, 2008; Fercho & Baugh, 2014; Kluzik et al., 2008; Kong et al., 2017; Wilke et al., 2013). In normal gravity, consistent mass-weight relationships lead to heavy reliance on weight-related cues for mass estimation (Bock, 1998; Bock et al., 1996b; Sangals et al., 1999). In microgravity, however, reduced weight cues become misleading and may systematically bias mass perception, affecting feedforward control of movement. We note, however, that the same proprioceptive degradation could also affect motor output through other pathways—for instance, by reducing tonic muscle activation or altering spinal reflex gains (Fisk et al., 1993)—independent of any explicit misrepresentation of body mass. Disentangling these mechanisms will require future experiments that can independently manipulate proprioceptive loading and limb inertia.

Earlier version: The persistence of deviant feedforward control and the subsequent within-movement corrections appears intriguing from the theories of human motor learning. Humans demonstrate remarkable ability to adapt to environmental changes in terrestrial studies. For experimentally-applied visual or force perturbations, individuals can recalibrate their sensorimotor control and restore baseline performance within minutes (Krakauer et al., 2019), a process often theorized as forming internal models of the new environment (Kawato, 1999; Wolpert et al., 1995). Relatedly, even when the body mass is perturbed by attaching a weight on the forearm, individuals can fully adapt their reaching movements within tens of trials (Wang & Sainburg, 2004). Previous studies on parabolic flights have also shown that the control of arm movements exhibited rapid adaptation to gravitational changes over the course of several parabolas. (e.g., Gaveau et al., 2016). In contrast, our in-flight sessions were scheduled at least three weeks post-launch—even well beyond the 2-3 weeks window typically considered critical for sensorimotor adaptation in spaceflight (Kanas & Manzey, 2008), but the participants still exhibited un-adapted feedforward control and internal model.

Revised version: The persistence of deviant feedforward control and the subsequent within-movement corrections appears intriguing from the theories of human motor learning. Humans demonstrate remarkable ability to adapt to environmental changes in terrestrial studies. For experimentally-applied visual or force perturbations, individuals can recalibrate their sensorimotor control and restore baseline performance within minutes (Krakauer et al., 2019), a process often theorized as forming internal models of the new environment (Kawato, 1999; Wolpert et al., 1995). Relatedly, even when the body mass is perturbed by attaching a weight on the forearm, individuals can fully adapt their reaching movements within tens of trials (Wang & Sainburg, 2004). Previous studies on parabolic flights have also shown that the control of arm movements exhibited rapid adaptation to gravitational changes over the course of several parabolas. (e.g., Gaveau et al., 2016). In contrast, our in-flight sessions were scheduled at least three weeks post-launch—even well beyond the 2-3 weeks window typically considered critical for sensorimotor adaptation in spaceflight (Kanas & Manzey, 2008), but the participants still exhibited un-adapted feedforward control and internal model, suggesting that the underlying sensory bias was not corrected through experience.

Earlier version: We believe this contrast stems from distinct natures of perturbation effects in microgravity when compared to in terrestrial environments. Previous studies on sensorimotor adaptation simulated environmental changes by imposing a perturbation to alter the relationship between motor commands and sensory consequences of motor commands. For reaching adaptation here, visuomotor perturbations changes the spatial mapping between hand motion and its visual representation’s motion (Krakauer et al., 2000; Wolpert et al., 2011); force perturbations change the dynamic mapping between actual force outputs and hand motion (Lackner & Dizio, 1994; Shadmehr & Mussa-Ivaldi, 1994). These mapping changes, though novel, are consistent thus the sensorimotor system is capable of approximating them and adapting accordingly. Theories of sensorimotor adaptation conceptualize the formation of the new mapping as acquiring internal models, which links motor command and efference copy to sensory consequences of motor commands (Flanagan & Wing, 1997; Kurtzer et al., 2008; Wolpert et al., 1998; Wolpert & Miall, 1996). However, the novel environment of microgravity alters the sensory estimate of motor apparatus (i.e., the body mass here), not merely the sensorimotor mapping. In normal gravity, weight-related sensory cues—mostly proprioceptive feedback—accurately represent body mass. In microgravity, however, these weight-related cues are substantially reduced due to the absence of gravitational pull, providing persistently biased information about body mass. In other words, the microgravity environment failed to provide veridical sensory cues for mass estimation, and reduced gravitational and proprioceptive cues persistently misinform the controller about body mass. Since body mass is a parametric input to the internal models, its bias would affect motor actions even when the internal models are well adapted. Hence, the unique environment of microgravity reveals a fundamental constraint of the sensorimotor system, i.e., its quick adaptability is restricted to learning a novel sensorimotor mapping, not to persistent sensory bias of bodily property.

Revised version: We believe this contrast stems from distinct natures of perturbation effects in microgravity when compared to in terrestrial environments. Previous studies on sensorimotor adaptation simulated environmental changes by imposing a perturbation to alter the relationship between motor commands and sensory consequences of motor commands. For reaching adaptation here, visuomotor perturbations changes the spatial mapping between hand motion and its visual representation’s motion (Krakauer et al., 2000; Wolpert et al., 2011); force perturbations change the dynamic mapping between actual force outputs and hand motion (Lackner & Dizio, 1994; Shadmehr & Mussa-Ivaldi, 1994). These mapping changes, though novel, are consistent thus the sensorimotor system is capable of approximating them and adapting accordingly. Theories of sensorimotor adaptation conceptualize the formation of the new mapping as acquiring internal models, which links motor command and efference copy to sensory consequences of motor commands (Flanagan & Wing, 1997; Kurtzer et al., 2008; Wolpert et al., 1998; Wolpert & Miall, 1996). However, the novel environment of microgravity may alter the sensory estimate of motor apparatus (i.e., the body mass here), not merely the sensorimotor mapping. In normal gravity, weight-related sensory cues—mostly proprioceptive feedback—accurately represent body mass. In microgravity, however, these weight-related cues are substantially reduced due to the absence of gravitational pull, providing persistently biased information about body mass. In other words, the microgravity environment failed to provide veridical sensory cues for mass estimation, and reduced gravitational and proprioceptive cues persistently misinform the controller about body mass. Since body mass is a parametric input to the internal models, its bias would affect motor actions even when the internal models are well adapted. Hence, the unique environment of microgravity may reveal an important constraint of the sensorimotor system, i.e., its quick adaptability is restricted to learning a novel sensorimotor mapping, not to persistent sensory bias of bodily property.

Earlier version: The discrete nature of our reaching task may contribute to the persistence of sensory bias. Within-movement corrections indicates that the sensorimotor system learns about sensory bias of mass estimation (Dimitriou et al., 2013), possibly via kinesthetic feedback during movement (Papaxanthis et al., 2005; Proske & Weber, 2023; Weber & Stelzer, 2022). However, movement-related cues only become available after movement initiation. Once a movement ends and the hand returns to rest, microgravity continues to elicit misleading sensory cues that bias mass estimation between trials. This lack of between-trial adaptation parallels findings from a manual interception study aboard the International Space Station, where astronauts persistently initiated movements too early when attempting to intercept an object moving at constant speed (McIntyre et al., 2001). Although astronauts corrected their movements online within each trial, their initial timing error persisted across multiple trials and sessions during a 15-day spaceflight. Hence, discrete motor tasks, such as manual interception and reaching, rely heavily on feedforward control, and are susceptible to sensory biases induced by microgravity. This explains why continuous reaching between targets, without stopping, shows quicker adaptation to zero gravity in parabolic flights (Papaxanthis et al., 2005) and does not exhibit prolonged movement duration during spaceflight (Fowler et al., 2008).

Revised version: The discrete nature of our reaching task may contribute to the persistence of sensory bias. Within-movement corrections indicates that the sensorimotor system learns about sensory bias of mass estimation (Dimitriou et al., 2013), possibly via kinesthetic feedback during movement (Papaxanthis et al., 2005; Proske & Weber, 2023; Weber & Stelzer, 2022). However, movement-related cues only become available after movement initiation. Once a movement ends and the hand returns to rest, microgravity continues to elicit misleading sensory cues that bias mass estimation between trials. This lack of between-trial adaptation parallels findings from a manual interception study aboard the International Space Station, where astronauts persistently initiated movements too early when attempting to intercept an object moving at constant speed (McIntyre et al., 2001). Although astronauts corrected their movements online within each trial, their initial timing error persisted across multiple trials and sessions during a 15-day spaceflight. Hence, discrete motor tasks, such as manual interception and reaching, rely heavily on feedforward control, and are susceptible to sensory biases induced by microgravity. This would also be consistent with the observation that continuous reaching between targets, without stopping, shows quicker adaptation to zero gravity in parabolic flights (Papaxanthis et al., 2005) and does not exhibit prolonged movement duration during spaceflight (Fowler et al., 2008).

Earlier version: What are the potential consequences of body mass underestimation beyond the simple reaching movements examined here? The implications may be particularly significant given two key considerations. First, the majority of human movements are not deliberately controlled (Brooks, 1986; Marsden et al., 1976), suggesting that underactuated initial movements may be pervasive in microgravity—especially for actions involving larger body masses such as whole-body movements or object manipulation. Without the controlled conditions and explicit speed requirements of our experimental task, such underactuation might manifest even more prominently in natural movements during spaceflight. Second, our finding of increased reliance on feedback-based corrections points to a potentially costly compensatory mechanism. As feedforward control becomes systematically biased by mass underestimation, the sensorimotor system must depend more heavily on feedback control to achieve adequate performance. This increased reliance on feedback processes may demand additional cognitive resources for action regulation (Bock et al., 2010), potentially contributing to the non-specific stressors that have been shown to impair human performance in space (Tian et al., 2024). The cognitive cost of these compensatory adjustments might be particularly relevant for complex motor tasks or situations requiring divided attention during spaceflight operations.

Revised version: What are the potential consequences of the underactuation observed in microgravity beyond the simple reaching movements examined here? The implications may be particularly significant given two key considerations. First, the majority of human movements are not deliberately controlled (Brooks, 1986; Marsden et al., 1976), suggesting that underactuated initial movements may be pervasive in microgravity—especially for actions involving larger body masses such as whole-body movements or object manipulation. Without the controlled conditions and explicit speed requirements of our experimental task, such underactuation might manifest even more prominently in natural movements during spaceflight. Second, our finding of increased reliance on feedback-based corrections points to a potentially costly compensatory mechanism. As feedforward control becomes systematically under-actuated, whether due to mass underestimation, altered neuromuscular properties, or both, the sensorimotor system must depend more heavily on feedback control to achieve adequate performance. This increased reliance on feedback processes may demand additional cognitive resources for action regulation (Bock et al., 2010), potentially contributing to the non-specific stressors that have been shown to impair human performance in space (Tian et al., 2024). The cognitive cost of these compensatory adjustments might be particularly relevant for complex motor tasks or situations requiring divided attention during spaceflight operations.

Earlier version: The preservation of speed-accuracy trade-offs in our study extends previous findings from Fitts’ task experiments performed during spaceflight (Fowler et al., 2008). Both the traditional speed-accuracy trade-off and action planning-dependent trade-off remained intact in microgravity (Figure 2—figure supplement 3 and 4), suggesting preserved fundamental motor control capacity. Motor performance during spaceflight typically degrades in tasks with high cognitive demands (Eddy et al., 1998; Manzey et al., 1998; Strangman et al., 2014), particularly those requiring sustained attention (Tian et al., 2024). Our findings suggest that while mass underestimation affects movement execution, the underlying motor control capabilities remain robust in microgravity when cognitive demands are modest.

Revised version: The preservation of speed-accuracy trade-offs in our study extends previous findings from Fitts’ task experiments performed during spaceflight (Fowler et al., 2008). Both the traditional speed-accuracy trade-off and action planning-dependent trade-off remained intact in microgravity (Figure 2—figure supplement 3 and 4), suggesting preserved fundamental motor control capacity. Motor performance during spaceflight typically degrades in tasks with high cognitive demands (Eddy et al., 1998; Manzey et al., 1998; Strangman et al., 2014), particularly those requiring sustained attention (Tian et al., 2024). Our findings suggest that while microgravity-induced underactuation affects movement execution, the underlying motor control capabilities remain robust in microgravity when cognitive demands are modest.

Earlier version: Our study presents three methodological limitations that warrant consideration. First, despite matching for age and gender, our ground controls exhibited faster movements than the taikonauts. This systematic difference, reflected in higher peak acceleration and speed, may stem from differing experience with the touchpad device used for movement recording. While taikonauts used touchpads regularly, 9 of 12 control participants had minimal experience, potentially leading to compensatory faster movements to avoid overtime errors. Second, the temporal evolution of mass underestimation effects remains unclear. Our earliest measurements occurred three weeks post-launch, though parabolic flight studies suggest these effects may emerge within hours of microgravity exposure (Crevecoeur et al., 2014; Papaxanthis et al., 2005). Additionally, incomplete post-flight recovery in some kinematic measures suggests either a mass overestimation due to exaggerated weight-related cues or neuromuscular deficiency typically observed shortly after returning to earth (Tays et al., 2021). Third, while our hand-reaching task enabled us to use well-established analytical frameworks for studying motor control aspects of movements, the generalizability of mass underestimation effects to other movements, particularly whole-body actions, remains to be determined. Future research should address these limitations by implementing matched device training protocols, adding early- and post-flight measurements, and extending investigations to diverse motor tasks. We also notice that the microgravity effect was less consistent for the 45° reaching direction. This direction engages more single-joint, elbow-dominant reaches, whereas the 90° and 135° targets require greater multi-joint (shoulder + elbow) coordination. Consistent with this, an exploratory analysis of hand-path curvature (cumulative curvature; positive = counterclockwise) showed larger curvature at 45° (6.484° ± 0.841°) than at 90° (1.539° ± 0.462°) or 135° (2.819° ± 0.538°). The significantly larger curvature in the 45° condition suggests that these movements deviate more from a straight-line path, supporting a more elbow-dominant movement in this direction. Importantly, this curvature pattern was present in both the pre-flight and in-flight phases, indicating that it is a general movement characteristic rather than a microgravity-induced effect. We postulate that single-joint movements, more so than we considered in our ideal two-link model, made effective mass rather small in the 45° direction, hence the mass underestimation and underactuation in feedforward planning had smaller effect, compared to model simulations, in this direction.

Revised version: Our study presents four methodological limitations that warrant consideration. First, despite matching for age and gender, our ground controls exhibited faster movements than the taikonauts. This systematic difference, reflected in higher peak acceleration and speed, may stem from differing experience with the touchpad device used for movement recording. While taikonauts used touchpads regularly, 9 of 12 control participants had minimal experience, potentially leading to compensatory faster movements to avoid overtime errors. Second, the temporal evolution of mass underestimation effects remains unclear. Our earliest measurements occurred three weeks post-launch, though parabolic flight studies suggest these effects may emerge within hours of microgravity exposure (Crevecoeur et al., 2014; Papaxanthis et al., 2005). Additionally, incomplete post-flight recovery in some kinematic measures suggests either a mass overestimation due to exaggerated weight-related cues or neuromuscular deficiency typically observed shortly after returning to earth (Tays et al., 2021). Third, while our hand-reaching task enabled us to use well-established analytical frameworks for studying motor control aspects of movements, the generalizability of mass underestimation effects to other movements, particularly whole-body actions, remains to be determined. Future research should address these limitations by implementing matched device training protocols, adding early-and post-flight measurements, and extending investigations to diverse motor tasks. We also notice that the microgravity effect was less consistent for the 45° reaching direction. This direction engages more single-joint, elbow-dominant reaches, whereas the 90° and 135° targets require greater multi-joint (shoulder + elbow) coordination. Consistent with this, an exploratory analysis of hand-path curvature (cumulative curvature; positive = counterclockwise) showed larger curvature at 45° (6.484° ± 0.841°) than at 90° (1.539° ± 0.462°) or 135° (2.819° ± 0.538°). The significantly larger curvature in the 45° condition suggests that these movements deviate more from a straight-line path, supporting a more elbow-dominant movement in this direction. Importantly, this curvature pattern was present in both the pre-flight and in-flight phases, indicating that it is a general movement characteristic rather than a microgravity-induced effect. We postulate that single-joint movements, more so than we considered in our ideal two-link model, made effective mass rather small in the 45° direction, hence the mass underestimation and underactuation in feedforward planning had smaller effect, compared to model simulations, in this direction. Fourth, our interpretation of the observed underactuation relies on a simplified optimal control model that treats muscles as ideal torque generators. This model does not capture several neurophysiological factors that may be altered in microgravity, including changes in tonic muscle activation, spinal reflex modulation, and descending vestibular influences on motor neuron excitability (Fisk et al., 1993), nor potential alterations in the damping and natural frequency of the limb’s mechanical response. However, the available evidence suggests that upper limb muscle capacity is relatively preserved in microgravity: a systematic review found that upper limb maximal voluntary contraction remained mostly unchanged during unloading periods of up to 45 days, with upper limb muscles declining substantially more slowly than lower limb and trunk muscles (Winnard et al., 2019; Bosutti et al., 2025). While the direction-dependent pattern of our results is most parsimoniously explained by a misrepresentation of inertial properties, we cannot exclude that some portion of the observed underactuation arises from these other neuromuscular factors. Distinguishing between mass underestimation and other sources of underactuation would benefit from future studies that combine detailed musculoskeletal modeling with direct measurements of muscle activation, joint impedance, and trunk kinematics during spaceflight.

Earlier version: In conclusion, our study provides compelling evidence that movement slowing in microgravity stems from systematic underestimation of body mass, rather than a conservative control strategy. By leveraging the direction-dependent variation in limb effective mass, we demonstrated that reaching movements in microgravity exhibit kinematic signatures—reduced and earlier peak speed/acceleration—that precisely match the predictions of mass underestimation. The persistence of these effects throughout long-duration spaceflight, coupled with increased feedback-based corrections, reveals how altered sensory inputs fundamentally impact both predictive and reactive components of motor control and evades sensorimotor adaptation. While basic motor control capabilities remain intact, as evidenced by preserved speed-accuracy trade-offs, the consistent underactuation of movements may have significant implications for the theorization of sensorimotor learning and for practical considerations of motor performance during space operations. These findings not only advance our understanding of how the sensorimotor system adapts to novel gravitational environments but also reveal fundamental principles of sensorimotor integration, particularly the central role of gravitational cues in shaping our perception of body dynamics.

Revised version: In conclusion, our study provides converging evidence that movement slowing in microgravity is associated with systematic underactuation of initial movements, a pattern most consistent with the underestimation of body mass hypothesis. By leveraging the direction-dependent variation in limb effective mass, we showed that reaching movements in microgravity exhibit kinematic signatures—reduced and earlier peak speed/acceleration—that are broadly consistent with the predictions of the mass underestimation hypothesis. Although other neurophysiological factors such as altered muscle tone and vestibular modulation may also contribute, these effects are not readily explained by a uniform strategic adjustment or general neuromuscular change alone. The persistence of these effects throughout long-duration spaceflight, coupled with increased feedback-based corrections, suggests how altered sensory inputs can impact both predictive and reactive components of motor control in ways that resist sensorimotor adaptation. While basic motor control capabilities remain intact, as evidenced by preserved speed-accuracy trade-offs, the consistent underactuation of movements may have significant implications for the theorization of sensorimotor learning and for practical considerations of motor performance during space operations. These findings not only advance our understanding of how the sensorimotor system adapts to novel gravitational environments but also highlight fundamental principles of sensorimotor integration, particularly the central role of gravitational cues in shaping our perception of body dynamics.

Methods

Participants

Twelve taikonauts (2 females, 10 males; mean age 49.5 ± 6.5 years) from the first four missions of the China Space Station (CSS) served as the experimental group (Table S1). The control group consisted of 12 right-handed residents from Beijing (2 females, 10 males; mean age 49.9 ± 6.5 years). Control group received monetary compensation. All participants were provided with written informed consent forms, approved by the Ethics Committee of the China Astronaut Research and Training Center and the Institutional Review Board of Peking University. The experimental group was exposed to microgravity for 92–187 days during manned missions in China’s Shenzhou Program. Due to operational constraints, the number of test sessions varied among taikonauts: one completed 4 sessions, five completed 6, and six completed 7 (Table S1). All control sessions were conducted at Peking University.

Task and experimental design

The task was conducted using a tablet (Surface Pro 6, Microsoft, Redmond, Washington) with 100Hz sampling frequency and noise-canceling earphones (FiiO FA7, FiiO Electronics Technology, Guangzhou, China). During in-flight sessions, participants wore earphones and assumed a neutral posture with feet secured by foot straps in front of a foldable tabletop, positioned 0.9 meters above the cabin floor. They held the tabletop edge with their left hand, allowing their right hand to move freely on the tablet (Figure 7A). The foot straps and left-hand grip provided body stabilization. The tablet, attached to the tabletop with Velcro, was placed ~35 cm in front of participants and 30 cm below chest level (Figure 7A). At the start of each trial, an orange circle (0.5 cm radius) appeared at the bottom of the screen, signaling participants to move their right index finger to this start circle. After a random delay (500–1000 ms), a target appeared at one of three possible locations, 12 cm from the start circle at 45°, 90°, or 135° counterclockwise from the horizontal axis (Figure 7A). Participants were instructed to reach the target quickly and accurately, pausing briefly before returning to the start. If the time from target appearance to movement end (finger speed <10 cm/s) exceeded 650 ms, a “too slow” message prompted faster movement. In half of the trials, a 50 ms “beep” sound was played at target appearance. Each session included 120 trials, with 3 target directions and 2 beep conditions balanced and randomized (Figure 7B). The x-y position of the index finger was recorded continuously at 100 Hz.

Experimental Setup and Design.

A) Top-down view of a participant performing the reaching task with the right hand on a tablet. The start position and all possible target locations are shown as orange dots on the tablet screen. B) Experimental design. Both groups completed 4–7 sessions, with each session consisting of 120 trials. Some taikonauts missed 1 or 2 in-flight sessions.

All experiments aboard the China Space Station were monitored in real time by experimenters at the Beijing Control Center. Task instructions were displayed on the opening screen of the data-acquisition application, and participants were given ample time to read them. The same team of experimenters administered all pre-, in-, and post-flight sessions using identical instructions. Consistent with standard practice, the astronauts served as both participants and on-orbit experimenters and were extensively trained for this role on the ground. Multiple pre-flight sessions were conducted to familiarize them with the task. These safeguards were implemented to ensure high-quality data.

The control group was included to assess a potential confounding effect of repeated measurements. The interval between successive test sessions was 6–9 days, which was shorter than those of the taikonauts (1–2 months) to conservatively evaluate potential practice effects. Despite the shorter intervals, no significant differences were observed across sessions for most measures (except reaction time), confirming minimal practice effects and providing a robust baseline for evaluating microgravity’s impact on the taikonauts. All ground tests, for both groups, were conducted using identical devices and software. The layout of the table and tablet was kept consistent across sessions. The only difference was that, in ground tests, participants sat in a chair approximately 90 cm above the ground without foot straps and did not need to hold the edge of the tabletop with their left hand for stability.

Model simulation

We simulated the peak speed, peak acceleration, and their corresponding times for reaching movements in different directions. We employed the classical two-joint arm model to demonstrate the effect of biomechanics (Li & Todorov, 2004; Figure 1B). This was combined with a movement utility model to estimate planned movement time (Shadmehr et al., 2016; Figure S6A and S5B) and a forward optimal controller to simulate reaching movements (Todorov, 2005; Figure 1C and F, Figure S6C-D). Our simulation is briefly outlined below in steps.

Earlier version: We used a two-joint planar arm model to estimate the effective mass of the hand moving in different directions (Figure 1B). The parameters of the upper arm (length l1, mass m1, center of mass lc1, inertia I1) and the lower arm (length l2, mass m2, center of mass lc2, inertia I2) are based on values from Shadmehr et al., 2016: The elbow and shoulder angles are denoted by θe and θs, respectively. The initial arm configuration is set by θs = 0. 785 rad, and θe = 1. 571 rad, following the experimental conditions from Gordon et al., 1982. With a particular joint angle configuration, the end-point hand position is: The arm’s inertia matrix is: where: with: The joint torques at the shoulder and elbow are represented as: and the forces at the hand are: We use the Jacobian matrix to relate force and torque, based on virtual work principle: Put S3 into S5: then combine S7 and S8: Since the hand speed at the beginning of the movement is zero, according to S4 and S9, the mass matrix of hand M(θ) is: The mass matrix M(θ) is a 2 x 2 matrix, and the effective mass m(θ) is determined by the length of the force vector when M(θ) is subjected to a unit acceleration.

Revised version: We used a two-joint planar arm model to estimate the effective mass of the hand moving in different directions (Figure 1B). The parameters of the upper arm (length l1, mass m1, center of mass lc1, inertia I1) and the lower arm (length l2, mass m2, center of mass lc2, inertia I2) are based on values from Shadmehr et al., 2016: l1=0.33l2=0.43mm1=1.93m2=1.52kglc1=l1/2lc2=2l2/3mI1=0.014I2=0.019kg The elbow and shoulder angles are denoted by θe and θs, respectively. The initial arm configuration is set by θs = 0.785 rad, and θe = 1.571 rad, following the experimental conditions from Gordon et al., 1982. With a particular joint angle configuration, the end-point hand position is: (eq.1) x=[l1cos⁡θs+l2cos⁡(θs+θe);l1sin⁡θs+l2sin⁡(θs+θe)]. The arm’s inertia matrix is: (eq.2) I(θs,θe)=[I11,I12;I21,I22], where: I11=a3+a1l12+a4+2a2l1cos⁡θe,I12=I12=a2+l1cos⁡θe+a4,I22=a4 with: a1=m2,a2=m2lc2,a3=m1lc12+i1,a4=m2lc22+i2 The joint torques at the shoulder and elbow are represented as: (eq.3) τ=I(θ)θ¨, and the forces at the hand are: (eq.4) f=M(θ)x¨. We use the Jacobian matrix Λ=dxdθ to relate force and torque, based on virtual work principle: (eq.5) τ=ΛTf, (eq.6) x˙=Λθ˙, (eq.7) x¨=Λ˙θ˙+Λθ¨. Put S3 into S5: (eq.8) f=Λ−1TI(θ)θ¨, then combine S7 and S8: (eq.9) f=Λ−1TI(θ)Λ−1(x¨−Λ˙θ˙). Since the hand speed at the beginning of the movement is zero, according to S4 and S9, the mass matrix of hand M(θ) is: (eq.10) M(θ)=Λ−1TI(θ)Λ−1, The mass matrix M(θ) is a 2 x 2 matrix, and the effective mass m(θ) is determined by the length of the force vector when M(θ) is subjected to a unit acceleration.

Earlier version: In Figure S6A, the solid curve illustrates the effective mass across various directions under normal gravity, while the dashed curve represents the effective mass with a hypothetical 30% ()mass underestimation in microgravity. The colored lines indicate the effective mass amplitudes for the three target directions used in our experiment: 45°, 90°, and 135°.

Revised version: In Figure S6A, the solid curve illustrates the effective mass across various directions under normal gravity, while the dashed curve represents the effective mass with a hypothetical 30% (m0g(θ)=0.3m(θ)) mass underestimation in microgravity. The colored lines indicate the effective mass amplitudes for the three target directions used in our experiment: 45°, 90°, and 135°.

Movement time determined by optimal utility

Earlier version: Following Shadmehr et al., 2016, the utility of an reaching action is given by: where α is the reward, m is mass, d is the distance to be moved, T is the planned duration of the movement, and a, b and i are scaling factors of the effort cost. γ represents a temporal discounting factor. The optimal movement time (Topt) is the time that maximizes the utility Jut: The parameter values used in the simulation of Figure S6B were taken from Shadmehr et al., 2016: α = 1000, a = 15, b = 100, i = 1. 1, γ = 1. The movement distance d is set as 0.12 m, the target distance in the experiment. Figure 1—figure supplement 1B illustrates how the utility curves change with planned movement time in three experimental directions (45°, 90°, and 135°) under normal gravity, with colored dots indicating the optimal movement durations.

Revised version: Following Shadmehr et al., 2016, the utility of an reaching action is given by: (eq.11) Jut=α−amT−bmdi/T21+γT where α is the reward, m is mass, d is the distance to be moved, T is the planned duration of the movement, and a, b and i are scaling factors of the effort cost. γ represents a temporal discounting factor. The optimal movement time (Topt) is the time that maximizes the utility Jut : (eq.12) Topt=argT⁡max(Jut). The parameter values used in the simulation of Figure S6B were taken from Shadmehr 2016: α = 1000, a = 15, b = 100, i = 1.1, γ = 1. The movement distance d is set as 0.12 m, the target distance in the experiment. Figure 1—figure supplement 1B illustrates how the utility curves change with planned movement time in three experimental directions (45°, 90°, and 135°) under normal gravity, with colored dots indicating the optimal movement durations.

Removed in the revised version: Movement kinematics simulated by optimal control theory

Earlier version: After setting the movement time Topt, the position, speed and acceleration over time can be simulated based on an optimal control model (Todorov, 2005). The arm’s control operates as a second-order low-pass filter: with the time constants τ1 = τ2 = 40ms.

Revised version: Movement kinematics simulated by optimal control theory. After setting the movement time Topt, the position, speed and acceleration over time can be simulated based on an optimal control model (Todorov, 2005). The arm’s control operates as a second-order low-pass filter: (eq.13) ut=τ1τ2f¨t+(τ1−τ2)f˙t+ft, with the time constants τ1 = τ2 = 40ms.

Earlier version: The system state vector is [x= p (t); p (t); f (t); g (t); p *], where p is position of hand, f is force acting on the hand, and g is an auxiliary state variable, p* is the target position. The control law is given by: The system dynamics are: With the sensory feedback carries information about position, velocity and force: ξt is signal-dependent noise ξk ~N(0, (σc ut) 2). ωt is sensory noise ωk ~N(0, Ωω), where We define the sampling interval Δt, and the number of steps is N = Topt /Δt, the discrete-time dynamics of the above system are: Which is transformed into the matrices as: The controller seeks to minimize the cost function J by optimizing the command u: Here, we used the penalty values from Todorov (2005) for each state and control variable:

Revised version: The system state vector is x=[p(t);p˙(t);f(t);g(t);p∗], where p is position of hand, f is force acting on the hand, and g is an auxiliary state variable, p* is the target position. The control law is given by: ut=τ1g˙t+gt,gt=τ2f˙t+ft The system dynamics are: (eq.14) xt+1=Axt+B(ut+ξt),yt=Hxt+ωt With the sensory feedback carries information about position, velocity and force: yt=[p(t);p˙(t);f(t)]+ωt ξt is signal-dependent noise ξk∼N(0,(σcut)2). ωt is sensory noise ωk~N(0, Ωω), where Ωω=diag(σp2,σv2,σf2) We define the sampling interval Δt, and the number of steps is N = Topt/Δt, the discrete-time dynamics of the above system are: p(t+Δ)=p(t)+p˙(t)Δp˙(t+Δ)=p˙(t)+f(t)Δ/mf(t+Δ)=f(t)(1−Δτ2)+g(t)Δ/τ2g(t+Δ)=g(t)(1−Δτ1)+u(t)(1+σcεt)Δ/τ1 Which is transformed into the matrices as: A=[1Δ00001Δ/m00001Δ/τ200001−Δ/τ1000001]H=[100000100000100]B=[0;0;0;Δ/τ1;0]. The controller seeks to minimize the cost function J by optimizing the command u: (eq.15) J=∑t(xtTQxt+utTRut),R=wu,Q=ppT+vvT+ffT,p=[wp;0;0;0;−wp],v=[0;wv;0;0;0],f=[0;0;wf;0;0] Here, we used the penalty values from Todorov (2005) for each state and control variable: wu=0.00001,wp=2,wv=0.2,wf=0.02,σc=0.5,Ωω=(0.5diag[0.02m;0.2ms;1N])2. We first estimated the time-varying gains {Lk, Kk} correspond to the feedforward mapping and the feedback correction gain, respectively. The control law can be expressed as: uk=ukff+ukfb=−Lkxknom−Kk(x^k−xknom), where uk is the control input, xknom is the nominal planned state, x^k is the estimated state, Lk is the feedforward (nominal) control associated with the planned trajectory, and Kk is the time-varying feedback gain that corrects deviations from the plan. To define the motor plan for comparison with behavior, we then simulate the deterministic open-loop trajectory by turning off noise and disabling feedback corrections, i.e., uk=ukff=−Lkxknom. We then obtain the peak speed, peak acceleration, and their respective times.

Removed in the revised version: . We first estimated the time-varying gains {Lk, Kk } correspond to the feedforward mapping and the feedback correction gain, respectively. The control law can be expressed as: , where uk is the control input, is the nominal planned state, is the estimated state, Lk is the feedforward (nominal) control associated with the planned trajectory, and Kk is the time-varying feedback gain that corrects deviations from the plan. To define the motor plan for comparison with behavior, we then simulate the deterministic open-loop trajectory by turning off noise and disabling feedback corrections, i.e., We then obtain the peak speed, peak acceleration, and their respective times.

In summary, we apply the movement utility theory to calculate the optimal movement time for a simplified two-joint arm moving in a 2D plane. This movement time is primarily determined by the effective mass and temporal discounting of reward. Next, we employ the optimal control theory to model the kinematics of hand movement, including its position, speed, and acceleration. This approach also enables us to estimate the peak speed and acceleration, as well as their respective times, for reaches performed toward different target directions. To simulate the effect of microgravity, we introduce a mass underestimation factor to account for the kinematic changes associated with this perceptual error.

Data analysis

Removed in the revised version: Kinematic analysis

Earlier version: A Butterworth low-pass filter with a 20 Hz cutoff frequency was applied to the positional data. Movement distance at each time point was computed as the Euclidean distance between the finger’s current position and the center of the start circle. Speed was estimated by performing linear regression over a moving window of five consecutive positional data points, and acceleration was calculated similarly from the speed data. Movement onset was defined as the first time point at which acceleration exceeded 50 cm/s2, while movement offset was the first time point at which acceleration fell below 50 cm/s2. Reaction time was defined as the interval between target onset and movement onset, and movement duration as the interval between movement onset and offset. Movement endpoint error was computed as the Euclidean distance between the target and the finger’s position at movement offset. Peak speed and peak acceleration were identified as the maximum speed and acceleration occurring between movement onset and offset, respectively. The times of peak speed and peak acceleration were measured as their respective time intervals from movement onset, and their relative times were calculated by dividing each by the movement duration.

Revised version: Kinematic analysis. A Butterworth low-pass filter with a 20 Hz cutoff frequency was applied to the positional data. Movement distance at each time point was computed as the Euclidean distance between the finger’s current position and the center of the start circle. Speed was estimated by performing linear regression over a moving window of five consecutive positional data points, and acceleration was calculated similarly from the speed data. Movement onset was defined as the first time point at which acceleration exceeded 50 cm/s2, while movement offset was the first time point at which acceleration fell below 50 cm/s2. Reaction time was defined as the interval between target onset and movement onset, and movement duration as the interval between movement onset and offset. Movement endpoint error was computed as the Euclidean distance between the target and the finger’s position at movement offset. Peak speed and peak acceleration were identified as the maximum speed and acceleration occurring between movement onset and offset, respectively. The times of peak speed and peak acceleration were measured as their respective time intervals from movement onset, and their relative times were calculated by dividing each by the movement duration.

Removed in the revised version: Submovement extraction

Earlier version: To separate the primary submovement from the secondary corrective submovement, the speed profile of each reaching movement was decomposed into a sum of submovements. Following the optimization algorithm described by Rohrer and Hogan (2006), the speed profile was fitted as a sum of support-bounded lognormal (LGNB) curves. Each LGNB curve is defined as: where T0 and T1 represent the start and end times of the movement, D is a scaling parameter, and the parameters μ and σ determine the skewness and kurtosis of the underlying lognormal function. These five parameters enable each LGNB submovement to vary across a wide range of possible shapes.

Revised version: Submovement extraction. To separate the primary submovement from the secondary corrective submovement, the speed profile of each reaching movement was decomposed into a sum of submovements. Following the optimization algorithm described by Rohrer and Hogan (2006), the speed profile was fitted as a sum of support-bounded lognormal (LGNB) curves. Each LGNB curve is defined as: B(t)=(T1−T0)Dσ2π(t−T0)(T1−t)exp⁡{(−12σ2)[ln⁡(t−T0T1−t)−μ]2} where T0 and T1 represent the start and end times of the movement, D is a scaling parameter, and the parameters μ and σ determine the skewness and kurtosis of the underlying lognormal function. These five parameters enable each LGNB submovement to vary across a wide range of possible shapes.

Earlier version: In theory, the number of submovements could be treated as a free parameter. However, to avoid overfitting, we limited each movement to two submovements, given the brevity of the movement (average duration approx. 350 ms). For each trial, we initially fit the data using both one and two LGNB curves and then determined the optimal number of submovements using the “greedy” method from the previous study (Rohrer & Hogan, 2006). Briefly, the fitting error for each movement with one and two LGNB curves was calculated as: where G(t) is the actual movement speed profile, and F(t) is the fitted speed profile (the sum of the submovements). If the fitting error for a single submovement was below a 2% threshold, or if adding a second submovement reduced the fitting error by less than 2%, a single submovement was considered adequate. Otherwise, two submovements were included. Trials with large fitting errors (>10%) were discarded, accounting for 3.60% and 4.83% of trials in the experimental and control groups, respectively. The percentage of trials with two submovements was calculated for each movement direction and phase. The time difference between the peaks of the two submovements was also computed, and set to zero if only one submovement was detected. Data fitting was performed using Matlab’s fmincon function (MathWorks, Natick, MA, USA), and the robustness of fitting was confirmed by testing 20 random sets of initial parameters.

Revised version: In theory, the number of submovements could be treated as a free parameter. However, to avoid overfitting, we limited each movement to two submovements, given the brevity of the movement (average duration approx. 350 ms). For each trial, we initially fit the data using both one and two LGNB curves and then determined the optimal number of submovements using the “greedy” method from the previous study (Rohrer & Hogan, 2006). Briefly, the fitting error for each movement with one and two LGNB curves was calculated as: ε=∫|F(t)−G(t)|dt∫|G(t)|dt where G(t) is the actual movement speed profile, and F(t) is the fitted speed profile (the sum of the submovements). If the fitting error for a single submovement was below a 2% threshold, or if adding a second submovement reduced the fitting error by less than 2%, a single submovement was considered adequate. Otherwise, two submovements were included. Trials with large fitting errors (>10%) were discarded, accounting for 3.60% and 4.83% of trials in the experimental and control groups, respectively. The percentage of trials with two submovements was calculated for each movement direction and phase. The time difference between the peaks of the two submovements was also computed, and set to zero if only one submovement was detected. Data fitting was performed using Matlab’s fmincon function (MathWorks, Natick, MA, USA), and the robustness of fitting was confirmed by testing 20 random sets of initial parameters.

Removed in the revised version: Statistical analysis

Earlier version: First, invalid trials were excluded from further analysis, with three possible causes: early initiation (RT < 150 ms, indicating guessing of the target direction), affecting 0.74% and 0.70% of trials in the experimental and control groups, respectively; late initiation (RT > 400 ms, indicating inattention), affecting 1.13% and 0%; and measurement failures, affecting 1.16% and 0.70%. The proportion of invalid trials did not differ significantly across testing phases (Friedman’s tests, all p > 0.1).

Revised version: Statistical analysis. First, invalid trials were excluded from further analysis, with three possible causes: early initiation (RT < 150 ms, indicating guessing of the target direction), affecting 0.74% and 0.70% of trials in the experimental and control groups, respectively; late initiation (RT > 400 ms, indicating inattention), affecting 1.13% and 0%; and measurement failures, affecting 1.16% and 0.70%. The proportion of invalid trials did not differ significantly across testing phases (Friedman’s tests, all p > 0.1).

Most dependent measures—including movement duration, peak acceleration, peak speed, time to peak acceleration/speed, relative time to peak acceleration/speed, the percentage of submovements, and inter-peak intervals—were analyzed using two-way repeated-measures ANOVAs with a 3 (phase) × 3 (movement direction) design. Reaction time was analyzed using three-way repeated-measures ANOVAs with a 3 (phase) × 3 (movement direction) × 2 (beep/no beep) design. Greenhouse-Geisser corrections were applied when sphericity was violated (Kirk, 1968). Normality was assessed with Shapiro-Wilk tests, and if more than 10% of tests for a dependent variable violated normality, the Aligned Rank Transform (ART) procedure was applied before conducting repeated-measures ANOVA. Tukey’s HSD tests were used for pairwise comparisons.

Earlier version: To examine the relationship between submovement and movement duration, a Linear Mixed Model (LMM) was used to analyze the association between changes in the inter-peak interval of submovements (ΔIPI) and changes in movement duration (ΔMD). The dependent variable yΔMD represents the mean change in movement duration across two adjacent phases (pre-to-in and post-to-in transitions) for each target direction. Target direction (45°, 90°, 135°) and phase transitions (pre-to-in and post-to-in) were treated as within-subject factors. The predictors included three fixed factors: ΔIPI, phase transition (xpha), and direction (xdri). An intercept term (upar) was included as a random factor to account for individual differences in movement duration. The resulting LMM is as follows: where β0 denotes the overall intercept and ε denotes a Gaussian residual error.

Revised version: To examine the relationship between submovement and movement duration, a Linear Mixed Model (LMM) was used to analyze the association between changes in the inter-peak interval of submovements (ΔIPI) and changes in movement duration (ΔMD). The dependent variable yΔMD represents the mean change in movement duration across two adjacent phases (pre-to-in and post-to-in transitions) for each target direction. Target direction (45°, 90°, 135°) and phase transitions (pre-to-in and post-to-in) were treated as within-subject factors. The predictors included three fixed factors: ΔIPI, phase transition (xpha), and direction (xdir). An intercept term (upar) was included as a random factor to account for individual differences in movement duration. The resulting LMM is as follows: yΔMD=β0+β1xΔIPI+β2xpha+β3xdir+Zupar+ε where β0 denotes the overall intercept and ε denotes a Gaussian residual error.

All statistical analyses were performed using MATLAB’s Statistics and Machine Learning Toolbox. The significance level was set at α = 0.05, and all tests were two-tailed unless otherwise specified.

For the directional dependent effects (Figure 6), we applied repeated measures correlation to quantify within-subject associations (Bakdash & Marusich, 2017). Each participant contributed three repeated measurements under different directions. The predictor variable was identical across conditions, and the outcome variables were the measured responses of interest. Repeated measures correlations were computed separately for each of the four metrics, controlling for between-subject differences by including subject-specific intercepts and estimating a single common slope across participants.